# MATH 360 Python Assignment 1

* See [Mathematical Python](https://patrickwalls.github.io/mathematicalpython/) for an introduction to Python and Jupyter
* See [Introduction to Mathematical Modelling](https://ubcmath.github.io/MATH360/) for more examples
* Write solutions in the cells with `YOUR CODE HERE`
* Do **not** import any packages (other than the standard packages in the cell below)
* Run the tests to verify your solutions
* There are **hidden tests** therefore your solutions may not be entirely correct even if they pass the tests below
* Submit your `.ipynb` notebook file to Canvas (download from Syzygy to your machine and upload to Canvas)

```python
import numpy as np
import scipy.linalg as la
import scipy.integrate as spi
import matplotlib.pyplot as plt
```

## Problem 1

See [Mathematical Python > Differential Equations](https://patrickwalls.github.io/mathematicalpython/differential-equations/first-order/) to review how to use the function `scipy.integrate.odeint` to approximate solutions of differential equations. Consider a first order scalar equation with an initial condition:

$$
y' = f(y,t) \ \ , \ y(t_0) = y_0
$$

Note that we represent the right-hand side of the equation as a function $f(y,t)$ where the unknown function $y(t)$ is the first parameter in the function. This is the order that `odeint` expects.

### Part 1a (3 marks)

Consider the equation $y' = t^2 - y^2$, $y(0) = 1$. Write a function called `f1a` which represents the right hand side of the equation.

<!-- grade_id: cell-11a9aecf5ee607fb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e0ffa93ddaa57c6e, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify f1a is a Python function which takes two input parameters. (1 mark)
assert callable(f1a) , "f1a should be a Python function."
assert f1a(1,2) > 0 , "f1a(1,2) should return a positive number."
print("Test 1: Success!")
```

<!-- grade_id: cell-f64e7afd731fc2b9, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify f1a(1,2) = 3 and f1a(3,-2) = -5. (1 mark)
assert abs(f1a(1,2) - 3) < 1e-14 , "Value f1a(1,2) should be 3."
assert abs(f1a(3,-2) + 5) < 1e-14 , "Value f1a(3,-2) should be -5."
print("Test 2: Success!")
```

<!-- grade_id: cell-cfd238dcace553f3, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify f1a returns the correct values. This cell contains hidden tests. (1 mark)
```

### Part 1b (3 marks)

Use the function `odeint` to approximate the solution of the equation $y' = t^2 - y^2$, $y(0) = 1$, over the interval $[0,2]$ using 101 equally spaced points from 0 to 2. Save the result as `y1b`.

<!-- grade_id: cell-02c22ebae72f5e8f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-f1ff5a4d9537f04c, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify y1b is a NumPy array of length 101. (1 mark)
assert isinstance(y1b,np.ndarray) , "y1b should be a NumPy array."
assert len(y1b) == 101 , "y1b should have length 101."
print("Test 1: Success!")
```

<!-- grade_id: cell-2ccbaf87b14ccf2c, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values at t=0 and t=2. (1 mark)
assert abs(y1b[0] - 1.0) < 1e-14 , "Value y1b[0] should be 1."
assert abs(y1b[-1] - 1.6794589163875036) < 1e-14 , "Value y1b[-1] should be approximately 1.6794589."
print("Test 2: Success!")
```

<!-- grade_id: cell-c4ac90ec20195d99, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify y1b has the correct values. This cell contains hidden tests. (1 mark)
```

Plot the result:

```python
x1b = np.linspace(0,2,101)
plt.plot(x1b,y1b)
plt.grid(True)
plt.show()
```

## Problem 2

Consider the second order differential equaton:

$$
t^2 x'' + t x' + (t^2 - 1)x = 0
$$

### Part 2a (3 marks)

Convert the equation into a first order system with new variables $u_0 = x$ and $u_1 = x'$. Write a Python function called `f2a` to represent the right hand side of the system

$$
\frac{d \mathbf{u}}{dt} = \mathbf{f}(\mathbf{u},t)
$$

where $\mathbf{u} = [u_0,u_1]$.

<!-- grade_id: cell-a771fa878ecf5b44, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-c9be9d59c074d5a6, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify f2a is a Python function which takes two input parameters and returns an array. (1 mark)
assert callable(f2a) , "f2a should be a Python function."
assert len(f2a([0,1],1)) == 2 , "f2a should return an array of length 2."
print("Test 1: Success!")
```

<!-- grade_id: cell-658c979e36ee6141, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values at x=1, x'=0 and t=1. (1 mark)
assert np.allclose(f2a([0,1],1),[1,-1]) , "f2a([0,1],1) should return array [1,-1]."
print("Test 2: Success!")
```

<!-- grade_id: cell-2d50ec06a78a0d50, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify f2a returns the correct values. This cell contains hidden tests. (1 mark)
```

### Part 2b (3 marks)

Use the function `odeint` to approximate the solution $x(t)$ of the equation

$$
t^2 x'' + t x' + (t^2 - 1)x = 0 , \ x(1) = 1 , \ x'(1) = 0
$$

over the interval $[1,11]$ using 501 equally spaced points from 1 to 11. Save the result as `x2b`.

<!-- grade_id: cell-b983260236227286, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-8d06de3a96b7ea17, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x2b is a NumPy array of length 501. (1 mark)
assert isinstance(x2b,np.ndarray) , "x2b should be a NumPy array."
assert len(x2b) == 501 , "x2b should have length 501."
print("Test 1: Success!")
```

<!-- grade_id: cell-ebbd224e2d8aa195, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values at t=1 and t=11. (1 mark)
assert abs(x2b[0] - 1.0) < 1e-14 , "Value x2b[0] should be 1."
assert abs(x2b[-1] + 0.32505727130105394) < 1e-14 , "Value x2b[-1] should be approximately -0.32505727130105394."
print("Test 2: Success!")
```

<!-- grade_id: cell-0630cd466c42d013, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify x2b has the correct values. This cell contains hidden tests. (1 mark)
```

Plot the result:

```python
t = np.linspace(1,11,501)
plt.plot(t,x2b), plt.xlim([0,11]), plt.grid(True)
plt.show()
```

## Problem 3

Consider the first order system of differential equations:

$$
\begin{align*}
x_0' &= x_0 x_1 - 3x_2 \\
x_1' &= -x_0 x_1 + 2x_2 \\
x_2' &= x_0 x_1 - 2.25x_2
\end{align*}
$$

### Part 3a

Write a Python function called `f3a` to represent the right hand side of the system

$$
\frac{d \mathbf{x}}{dt} = \mathbf{f}(\mathbf{x},t)
$$

where $\mathbf{x} = [x_0,x_1,x_2]$.

<!-- grade_id: cell-695cbfb2bb03d442, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b0146e143dfa0ac0, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify f3a is a Python function which takes two input parameters and returns an array. (1 mark)
assert callable(f3a) , "f2a should be a Python function."
assert len(f3a([1,1,1],1)) == 3 , "f3a should return an array of length 3."
print("Test 1: Success!")
```

<!-- grade_id: cell-b4d6d42efc726798, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values at x0=x1=x2=1 and t=1. (1 mark)
assert np.allclose(f3a([1,1,1],1),[-2,1,-1.25]) , "f3a([1,1,1],1) should return the array [-2,1,-1.25]."
print("Test 2: Success!")
```

<!-- grade_id: cell-93d885f6a2753b94, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify f3a returns the correct values. This cell contains hidden tests. (1 mark)
```

### Part 3b (3 marks)

Use the function `odeint` to approximate the solution $\mathbf{x}(t)$ of the system above with initial values $x_0(0) = x_1(0) = x_2(0) = 1$. Approximate the solution over the interval $[0,10]$ using 501 equally spaced points from 0 to 10. Save the result as `x3b` (a NumPy array of size 501 by 3).

<!-- grade_id: cell-24f7e63781603bde, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-ba10ba53ba3e49a3, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x3b is a NumPy array of size 501 by 3. (1 mark)
assert isinstance(x3b,np.ndarray) , "x3b should be a NumPy array."
assert x3b.shape == (501,3) , "x3b should have shape (501,3)."
print("Test 1: Success!")
```

<!-- grade_id: cell-ebb55258af8b73f9, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values at t=0 and t=10. (1 mark)
assert np.allclose(x3b[0,:],[1,1,1]) , "Value x3b[0,:] should be the intial values [1,1,1]."
assert np.allclose(x3b[-1,:],[0.39684616, 1.76721871, 0.27379751]) , "Value x2b[-1,:] should be approximately [0.40, 1.77, 0.27]."
print("Test 2: Success!")
```

<!-- grade_id: cell-a50cb7ff86a738e3, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify x3b has the correct values. This cell contains hidden tests. (1 mark)
```

Plot the results:

```python
t = np.linspace(0,10,501)
plt.plot(t,x3b), plt.grid(True)
plt.show()
```

## Problem 4 (3 marks)

Consider the system of equations

$$
\begin{align*}
u_0' &= u_0 u_1 - u_1^2 \\
u_1' &= u_0 u_1 - c u_0^2
\end{align*}
$$

where $c$ is a parameter such that $c \ne 1$. There is only one steady state solution $(0,0)$ for any $c \ne 1$ and the Jacobian is the zero matrix therefore we cannot use linearization to classify the steady state. Plot several trajectories for different initial conditions and different values of the parameter $c$. Find values $a$ and $b$ defining an interval such that if $c \in (a,b)$ then $u_0(t) \to 0$ and $u_1(t) \to 0$ as $t \to \infty$ for any initial conditions. Save the values $a$ and $b$ as `a4` and `b4` respectively.

<!-- grade_id: cell-370883285ba0dcfd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a3a57578cf1ba569, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify a4 and b4 are numeric types. (1 mark)
assert isinstance(a4,(int,float))
assert isinstance(b4,(int,float))
print("Test 1: Success!")
```

<!-- grade_id: cell-f52af7733b28c64a, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify a4 is the correct value. This cell contains hidden tests. (1 mark)
```

<!-- grade_id: cell-a4596d3e5263eae3, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify b4 is the correct value. This cell contains hidden tests. (1 mark)
```