# MATH 360 Python Assignment 0

* See [Mathematical Python](https://patrickwalls.github.io/mathematicalpython/) for an introduction to Python and Jupyter
* Write solutions in the cells with `YOUR CODE HERE`
* Do **not** import any packages (other than the standard packages in the cell below)
* Run the tests to verify your solutions
* There are **hidden tests** therefore your solutions may not be entirely correct even if they pass the tests below
* Submit your `.ipynb` notebook file to Canvas (download from Syzygy to your machine and upload to Canvas)

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
```

## Problem 1: Numbers

See [Mathematical Python > Numbers](https://patrickwalls.github.io/mathematicalpython/python/numbers/) and [Mathematical Python > Variables](https://patrickwalls.github.io/mathematicalpython/python/variables/) to review how to perform numerical computations and save values to variables.

### Part 1a (2 marks)

Compute the following value for $r = 0.25$ and $N = 10$.

$$
\frac{1 - r^{N + 1}}{1 - r}
$$

Save the result as `x1a`.

<!-- grade_id: cell-52371be020c383a8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b440b2378931201a, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x1a is a floating point number larger than 0. (1 mark)
assert isinstance(x1a,float) , "x1a should be a float."
assert x1a > 0 , "x1a should be positive."
print("Test 1: Success!")
```

<!-- grade_id: cell-913584447aa1751c, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify x1a is the correct value. This cell contains hidden tests. (1 mark)
```

### Part 1b (2 marks)

Compute the following value for $a = 2$, $b = -3$ and $c = -1$.

$$
\frac{- b + \sqrt{b^2 - 4ac}}{2a}
$$

Save the result as `x1b`.

<!-- grade_id: cell-17626fb52c78fd2d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-8c0deb76311d136f, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x1b is a floating point number larger than 0. (1 mark)
assert isinstance(x1b,float) , "x1b should be a float."
assert x1b > 0 , "x1b should be positive."
print("Test 1: Success!")
```

<!-- grade_id: cell-cd817a8d8b4b31fa, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify x1b is the correct value. This cell contains hidden tests. (1 mark)
```

### Part 1c (2 marks)

Compute the following value for $a_0 = 3$, $a_1 = 7$, $a_2 = 15$, $a_3 = 1$ and $a_4 = 292$.

$$
a_0 + \frac{1}{a_1 + \frac{1}{a_2 + \frac{1}{a_3 + \frac{1}{a_4}}}}
$$

Save the result as `x1c`.

<!-- grade_id: cell-b59459e3690c59b6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6cefbaa36a234ca0, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x1c is a floating point number larger than 0. (1 mark)
assert isinstance(x1c,float) , "x1c should be a float."
assert x1c > 0 , "x1c should be positive."
print("Test 1: Success!")
```

<!-- grade_id: cell-93bfb86cdc805f26, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify x1c is the correct value. This cell contains hidden tests. (1 mark)
```

## Problem 2: Vectors

See [Mathematical Python > NumPy](https://patrickwalls.github.io/mathematicalpython/scipy/numpy/) to review how to work with NumPy arrays.

### Part 2a (2 marks)

Create a NumPy array with 121 evenly spaced values from 0 to 15 (inclusive). Save the result as `v2a`.

<!-- grade_id: cell-a9bcae4c88206fc9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6eaa9fe6031fcd84, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify v2a is a NumPy array of length 121 with first entry 0 and last entry 15. (1 mark)
assert isinstance(v2a,np.ndarray) , "v2a should be a NumPy array."
assert len(v2a) == 121 , "v2a should have length 121."
assert v2a[0] == 0. , "First entry of v2a should be 0."
assert v2a[-1] == 15. , "Last entry of v2a should be 15."
print("Test 1: Success!")
```

<!-- grade_id: cell-7530f160f0aa2b2e, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify v2a is the correct value. This cell contains hidden tests. (1 mark)
```

### Part 2b (2 marks)

Create a NumPy array with 21 entries such that entry at index $n$ is given by the formula:

$$
\frac{(-1)^n x^{n+1}}{n + 1}
$$

for $x = 1/2$. Save the result as `v2b`.

<!-- grade_id: cell-98e654a5c746ab9f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e3cd5b1e423f6e1d, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify v2b is a NumPy array of length 21 with first entry 1/2. (1 mark)
assert isinstance(v2b,np.ndarray) , "v2b should be a NumPy array."
assert len(v2b) == 21 , "v2b should have length 21."
assert abs(v2b[0] - 1/2) < 1e-14 , "First entry of v2b should be 1/2."
print("Test 1: Success!")
```

<!-- grade_id: cell-759f04b8f7313c58, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify v2b has correct values. This cell contains hidden tests. (1 mark)
```

### Part 2c (2 marks)

Use `v2b` to compute the sum:

$$
\sum_{n=1}^{21} \frac{(-1)^{n+1} x^n}{n}
$$

for $x = 1/2$. Save the result as `y2c`.

<!-- grade_id: cell-b19c860d1e6f62a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-d2bf9dcd3e0042b0, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify y2c is a float and is positive. (1 mark)
assert isinstance(y2c,float) , "y2c should be a float."
assert y2c > 0
print("Test 1: Success!")
```

<!-- grade_id: cell-fc3a0543b0086512, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify y2c is the correct value. This cell contains hidden tests. (1 mark)
```

## Problem 3: Matrices

See [Mathematical Python > Linear Algebra](https://patrickwalls.github.io/mathematicalpython/linear-algebra/linear-algebra-scipy/) to review how to solve linear systems of equations.

### Part 3a (2 marks)

Create the 40 by 40 square matrix $A$ with -2 along the main diagonal and +1 along the upper and lower diagonals and 0 everywhere else:

$$
A =
\begin{bmatrix}
-2 & +1 & 0 & & & \cdots & 0 \\
+1 & -2 & +1 & 0 & & \cdots & 0 \\
0 & +1 & -2 & +1 & 0 & \cdots & 0 \\
\vdots & & & \ddots & & & \vdots \\
0 & \dots & 0 & +1 & -2 & +1 & 0 \\
0 & \dots & & 0 & +1 & -2 & +1 \\
0 & \dots & & & 0 & +1 & -2
\end{bmatrix}
$$

Save the result as `A3a`.

<!-- grade_id: cell-f39bb2582344f08f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b1623ab737cb6dc8, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify A3a is a NumPy array of size 40 by 40. (1 mark)
assert isinstance(A3a,np.ndarray) , "A3a should be a NumPy array."
assert np.shape(A3a) == (40,40) , "A3a should have shape 40 by 40."
assert np.allclose(A3a[1,:3],[1,-2,1]) , "Row at index 1 should begin with 1,-2,1."
print("Test 1: Success!")
```

<!-- grade_id: cell-25636334d51810a3, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify A3a has the correct values. This cell contains hidden tests. (1 mark)
```

### Part 3b (2 marks)

Compute the solution $\mathbf{x}$ of the system $A \mathbf{x} = \mathbf{b}$ where $A$ is the matrix from part 3(b) above and $\mathbf{b}$ is a vector of all 1s. Save the result as `x3b`.

<!-- grade_id: cell-233c207717c8d04a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

Visualize the result. The plots should look like a parabola since $A$ is like a second derivative operator.

```python
plt.plot(x3b,'.-')
plt.show()
```

<!-- grade_id: cell-86de57b25b3d0f74, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x3b is a NumPy array of length 40. (1 mark)
assert isinstance(x3b,np.ndarray) , "x3b should be a NumPy array."
assert np.size(x3b) == 40 , "x3b should have length 40."
assert np.allclose(x3b[20],-210)
print("Test 1: Success!")
```

<!-- grade_id: cell-58557f5901ade802, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify x3b has the correct values. This cell contains hidden tests. (1 mark)
```

## Problem 4: Functions (5 marks)

See [Mathematical Python > Functions](https://patrickwalls.github.io/mathematicalpython/python/functions/) to review how to define Python functions.

### Part 4a (3 marks)

Write a function called `f4a` which takes input parameter `x` and returns the value

$$
\frac{1}{1 + x^2}
$$

<!-- grade_id: cell-abc55e42286380bb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-404a8246d514f128, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify f4a is a Python function takes a single input parameter and outputs a positive number. (1 mark)
assert callable(f4a) , "f4a should be a Python function."
assert f4a(0.1234) > 0 , "f4a should return a positive number."
print("Test 1: Success!")
```

<!-- grade_id: cell-c932f17a02cdc6b1, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify f4a(0) = 1 and f4a(-2) = 0.2. (1 mark)
assert abs(f4a(0) - 1.0) < 1e-14 , "Value f4a(0) should be 1.0."
assert abs(f4a(-2) - 0.2) < 1e-14 , "Value f4a(-2) should be 0.2."
print("Test 2: Success!")
```

<!-- grade_id: cell-a147ff806e96f9e1, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify f4a returns the correct values. This cell contains hidden tests. (1 mark)
```

### Part 4b (3 marks)

Write a function called `f4b` which takes input parameters `x` and `N` and returns the sum:

$$
\sum_{n=1}^{N} \frac{\cos(\pi n x)}{n}
$$

<!-- grade_id: cell-a2854595e2a4b5f2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-4b33f5652d195305, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify f4b is a Python function takes two input parameters and returns a number. (1 mark)
assert callable(f4b) , "f4b should be a Python function."
assert isinstance(f4b(0.12345,5),float) , "f4b should return a float."
print("Test 1: Success!")
```

<!-- grade_id: cell-a615d29dbf4c9c6b, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify f4a(0,2) = 1 + 1/2 and f4a(1,3) = -1 + 1/2 - 1/3. (1 mark)
assert abs(f4b(0,2) - 1.5) < 1e-14 , "Value f4a(0) should be 1.0."
assert abs(f4b(1,3) + 5/6) < 1e-14 , "Value f4a(-2) should be 0.2."
print("Test 2: Success!")
```

<!-- grade_id: cell-9733a98f01f6712f, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify f4a returns the correct values. This cell contains hidden tests. (1 mark)
```

## Problem 5: Plotting (3 marks)

See [Mathematical Python > Matplotlib](https://patrickwalls.github.io/mathematicalpython/scipy/matplotlib/) to review how to plot a function.

Plot the function $f(x) = e^{-x^2} \cos(10x)$ on the interval $[-2,2]$. Use 501 equally spaced $x$ values from -2 to 2 (inclusive). Save the vector of $x$ values as `x5` and the $y$ values as `y5`.

<!-- grade_id: cell-44f733b83c38419c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


plt.plot(x5,y5)
plt.show()
```

<!-- grade_id: cell-7a520b93e694242e, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify x5 is a NumPy array of length 501 with first entry -2 and last entry 2. (1 mark)
assert isinstance(x5,np.ndarray) , "x5 should be a NumPy array."
assert len(x5) == 501 , "x5 should have length 501."
assert x5[0] == -2. , "First entry of x5 should be -2."
assert x5[-1] == 2. , "Last entry of x5 should be 2."
print("Test 1: Success!")
```

<!-- grade_id: cell-63f1f7cebda873d9, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify y5 is a NumPy array of length 501 with entry 1 at index 250. (1 mark)
assert isinstance(y5,np.ndarray) , "y5 should be a NumPy array."
assert len(y5) == 501 , "y5 should have length 501."
assert y5[250] == 1. , "Entry of y5a at index 250 should be 1.0."
print("Test 2: Success!")
```

<!-- grade_id: cell-650684b4d136dfa3, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify y5 has the correct values. This cell contains hidden tests. (1 mark)
```