# MATH 307 Assignment 2

* See [Mathematical Python](https://patrickwalls.github.io/mathematicalpython/) for an introduction to Python and Jupyter
* See [MATH 307 Jupyter Notebooks](https://ubcmath.github.io/MATH307/) for examples
* Run the tests to verify your solutions
* There are **hidden tests** therefore your solutions may not be entirely correct even if they pass the tests below
* Submit your `.ipynb` notebook file to Canvas

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
import matplotlib as mpl
mpl.rc('image', cmap='RdBu')
```

## Problem 1 (6 marks)

A [Toeplitz matrix](https://en.wikipedia.org/wiki/Toeplitz_matrix) has constant values along the diagonals. The function `scipy.linalg.toeplitz` creates a Toeplitz matrix from a given vector. For example:

```python
M = la.toeplitz([-2,1,0,0,0,0,0])
print(M)
```

* Use the function `numpy.linalg.matrix_rank` to compute the rank of a matrix
* The nullspace of $M$ is denoted $N(M)$
* The range of $M$ is denoted $R(M)$

### Problem 1a (2 marks)

Compute the dimensions of $N(M_1)$ and $R(M_1)$ for the matrix $M_1$ below. Save responses as `nullspace_M1` and `range_M1` respectively.

```python
M1 = la.toeplitz([1,1,0,0,0,0,0,1])
print(M1)
```

<!-- grade_id: cell-47c828aad9d9bb00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
nullspace_M1 = 
range_M1 = 
```

<!-- grade_id: cell-3a4a194ab54daa87, points: 1, grade: true, solution: false -->
```python
"Verify nullspace_M1 and range_M1 are non-negative numbers. (1 mark)"
assert nullspace_M1 >= 0 and range_M1 >= 0
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-cf424c49ee898a9b, points: 1, grade: true, solution: false -->
```python
"Verify values of nullspace_M1 and range_M1. This cell contains hidden tests. (1 mark)"
assert nullspace_M1 >= 0 and range_M1 >= 0
print("Problem 1a Test 2: Success!")
```

### Problem 1b (2 marks)

Compute the dimensions of $N(M_2)$ and $R(M_2)$ for the matrix $M_2$ below. Save responses as `nullspace_M2` and `range_M2` respectively.

```python
M2 = la.toeplitz([2,1,1,0,0,0,1,1])
print(M2)
```

<!-- grade_id: cell-6c82573446f0d349, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
nullspace_M2 = 
range_M2 = 
```

<!-- grade_id: cell-2c9cb136a3c6d37c, points: 1, grade: true, solution: false -->
```python
"Verify nullspace_M2 and range_M2 are non-negative numbers. (1 mark)"
assert nullspace_M2 >= 0 and range_M2 >= 0
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-56db7cdb7ebbf6ed, points: 1, grade: true, solution: false -->
```python
"Verify values of nullspace_M2 and range_M2. This cell contains hidden tests. (1 mark)"
assert nullspace_M2 >= 0 and range_M2 >= 0
print("Problem 1b Test 2: Success!")
```

### Problem 1c (2 marks)

Compute the dimensions of $N(M_3)$ and $R(M_3)$ for the matrix $M_3$ below. Save responses as `nullspace_M3` and `range_M3` respectively.

```python
M3 = la.toeplitz([2,2,1,0,0,0,1,2])
print(M3)
```

<!-- grade_id: cell-028eb265830a1c9c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
nullspace_M3 = 
range_M3 = 
```

<!-- grade_id: cell-df81e64cd6e8994b, points: 1, grade: true, solution: false -->
```python
"Verify nullspace_M3 and range_M3 are non-negative numbers. (1 mark)"
assert nullspace_M3 >= 0 and range_M3 >= 0
print("Problem 1c Test 1: Success!")
```

<!-- grade_id: cell-591449a982f85e79, points: 1, grade: true, solution: false -->
```python
"Verify values of nullspace_M3 and range_M3. This cell contains hidden tests. (1 mark)"
assert nullspace_M3 >= 0 and range_M3 >= 0
print("Problem 1c Test 2: Success!")
```

## Problem 2 (6 marks)

The code cells below perform the following tasks:

* Create matrices $A_0,\dots,A_9$ each of size 5 x 3 which represent the digits 0 to 9
* Visualize the matrices as 5 x 3 pixel images
* Flatten each matrix $A_k$ into a vector $\mathbf{v}_k$ of length 15 and construct a 15 x 10 matrix $A$ with vectors $\mathbf{v}_0,\dots,\mathbf{v}_9$ in the columns:

$$
A = \begin{bmatrix} & & \\ \mathbf{v}_0 & \cdots & \mathbf{v}_9 \\ & & \end{bmatrix}
$$

Use the matrix $A$ to answer the problems (a)-(e) below.

```python
A0 = np.array([[1.,1.,1.],
               [1.,0.,1.],
               [1.,0.,1.],
               [1.,0.,1.],
               [1.,1.,1.]])

A1 = np.array([[0.,1.,0.],
               [0.,1.,0.],
               [0.,1.,0.],
               [0.,1.,0.],
               [0.,1.,0.]])

A2 = np.array([[1.,1.,1.],
               [0.,0.,1.],
               [1.,1.,1.],
               [1.,0.,0.],
               [1.,1.,1.]])

A3 = np.array([[1.,1.,1.],
               [0.,0.,1.],
               [0.,1.,1.],
               [0.,0.,1.],
               [1.,1.,1.]])

A4 = np.array([[1.,0.,1.],
               [1.,0.,1.],
               [1.,1.,1.],
               [0.,0.,1.],
               [0.,0.,1.]])

A5 = np.array([[1.,1.,1.],
               [1.,0.,0.],
               [1.,1.,1.],
               [0.,0.,1.],
               [1.,1.,1.]])

A6 = np.array([[1.,1.,1.],
               [1.,0.,0.],
               [1.,1.,1.],
               [1.,0.,1.],
               [1.,1.,1.]])

A7 = np.array([[1.,1.,1.],
               [0.,0.,1.],
               [0.,0.,1.],
               [0.,0.,1.],
               [0.,0.,1.]])

A8 = np.array([[1.,1.,1.],
               [1.,0.,1.],
               [1.,1.,1.],
               [1.,0.,1.],
               [1.,1.,1.]])

A9 = np.array([[1.,1.,1.],
               [1.,0.,1.],
               [1.,1.,1.],
               [0.,0.,1.],
               [1.,1.,1.]])

digits = [A0,A1,A2,A3,A4,A5,A6,A7,A8,A9]
plt.figure(figsize=(16,7))
for n,digit in enumerate(digits):
    plt.subplot(2,5,n+1)
    plt.imshow(digit)
    plt.title('$A_{}$'.format(n)), plt.colorbar(), plt.clim(-1,1)
plt.show()
```

Flatten each matrix into a vector of length 15 and put the vectors into the columns of a matrix $A$:

```python
A = np.column_stack([digit.flatten() for digit in digits])
A
```

Let $\mathbf{v}_k$ be the $k$th column of $A$. Remember that indices start at 0 in Python. The vector $\mathbf{v}_k$ reshaped into a $5 \times 3$ matrix is exactly $A_k$. For example, the column at index 4 in $A$ reshaped into a $5 \times 3$ matrix is displayed below.

```python
k = 4
vk = A[:,k]
plt.imshow(vk.reshape(5,3)), plt.colorbar(), plt.clim([-1,1])
plt.show()
```

### Problem 2a (1 mark)

Determine the rank of $A$ and save the value as `rankA`.

<!-- grade_id: cell-f38cdbea574b182f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
rankA = 
```

<!-- grade_id: cell-ae51ea4612d01a55, points: 1, grade: true, solution: false -->
```python
"Verify value of rankA. This cell contains hidden tests. (1 mark)"
assert rankA > 0
print("Problem 2a Test 1: Success!")
```

### Problem 2b (1 mark)

**True** or **False**: The vectors $\mathbf{v}_0,\dots,\mathbf{v}_9$ are linearly independent. Save your response as `observation2b`.

<!-- grade_id: cell-972e07484c2bbe9a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
observation2b = 
```

<!-- grade_id: cell-609ce3f5c86a8c6d, points: 1, grade: true, solution: false -->
```python
"Verify value of observation2b. This cell contains hidden tests. (1 mark)"
assert isinstance(observation2b,bool)
print("Problem 2b Test 1: Success!")
```

### Problem 2c (2 marks)

The digit 8 looks like a combination of 2 and 5. Compute the orthogonal projection of $\mathbf{v}_8$ onto $\mathrm{span} \{ \mathbf{v}_2 , \mathbf{v}_5 \}$. Save the result as `proj825`.

Note that you can select columns 2 and 5 from matrix `A` and save to variable `A25` with the command:

```
A25 = A[:,[2,5]]
```

and compute the thin QR decomposition using the function [`scipy.linalg.qr`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.qr.html) with the parameter `mode='economic'`:

```
Q1,R1 = la.qr(A25,mode='economic')
```

Recall the projection of a vector $\mathbf{v}$ onto $R(M)$ for some matrix $M$ is $Q_1Q_1^T\mathbf{v}$ where $M=Q_1R_1$.

<!-- grade_id: cell-01ed9ddf6ca08841, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
proj825 = 
```

<!-- grade_id: cell-8eda0bdc0404c991, points: 1, grade: true, solution: false -->
```python
"Verify type and size of proj825. (1 mark)"
assert isinstance(proj825,np.ndarray)
assert proj825.size == 15
print("Problem 2c Test 1: Success!")
```

<!-- grade_id: cell-430f4a4c312c5102, points: 1, grade: true, solution: false -->
```python
"Verify value of proj825. This cell contains hidden tests. (1 mark)"
assert np.allclose(proj825[:3].flatten(),[1.1 , 1.1 , 1.1])
print("Problem 2c Test 2: Success!")
```

Visualize the projection as a 5 x 3 pixel image:

```python
plt.imshow(proj825.reshape(5,3)), plt.colorbar(), plt.clim([-1,1])
plt.show()
```

### Problem 2d (2 marks)

Find the shortest distance from $\mathbf{v}_9$ to $\mathrm{span} \{ \mathbf{v}_0 , \mathbf{v}_1 , \mathbf{v}_2 , \mathbf{v}_3 , \mathbf{v}_4 \}$. Save the result as `distance9`. Use the function [`scipy.linalg.norm`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.norm.html) to compute the norm.

<!-- grade_id: cell-d78e2e61f1545d28, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
distance9 = 
```

<!-- grade_id: cell-ea07f610840230fa, points: 1, grade: true, solution: false -->
```python
"Verify distance9 is a positive number. (1 mark)"
assert distance9 > 0
print("Problem 2d Test 1: Success!")
```

<!-- grade_id: cell-fc8fe1d419a84b5a, points: 1, grade: true, solution: false -->
```python
"Verify value of distance9. This cell contains hidden tests. (1 mark)"
assert np.round(distance9,1) == 0.7
print("Problem 2d Test 2: Success!")
```

## Problem 3 (8 marks)

Let $M = Q_1 R_1$ be the thin QR decomposition of $M$. The projection of a vector $\mathbf{v}$ onto the range $R(M)$ is given by

$$
\mathrm{proj}_{R(M)} (\mathbf{v}) = Q_1 Q_1^T \mathbf{v}
$$

Compute the thin QR decomposition using the function [`scipy.linalg.qr`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.qr.html) with the parameter `mode='economic'`:

```
Q1,R1 = la.qr(M,mode='economic')
```

Consider the matrix

$$
M = \begin{bmatrix} 1 & 1 & 0 \\ 1 & 1 & 1 \\ 0 & 1 & 1 \\ 1 & 0 & 1 \\ 1 & 2 & 1 \end{bmatrix}
$$

```python
M = np.array([[1.,1.,0.],
              [1.,1.,1.],
              [0.,1.,1.],
              [1.,0.,1.],
              [1.,2.,1.]])
```

### Problem 3a (2 marks)

Compute the thin QR decomposition $M = Q_1 R_1$. Save the result as `Q1M` and `R1M`.

<!-- grade_id: cell-eccf4f368340019b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
Q1M = 
R1M = 
```

<!-- grade_id: cell-983bb638c385a4c0, points: 1, grade: true, solution: false -->
```python
"Verify type of Q1M and R1M. (1 mark)"
assert isinstance(Q1M,np.ndarray)
assert isinstance(R1M,np.ndarray)
print("Problem 3a Test 1: Success!")
```

<!-- grade_id: cell-765f4cd67182ab93, points: 1, grade: true, solution: false -->
```python
"Verify values of Q1M and R1M. This cell contains hidden tests. (1 mark)"
assert np.allclose(Q1M[0,0],-0.5)
assert np.allclose(R1M[0,0],-2)
print("Problem 3a Test 2: Success!")
```

### Problem 3b (2 marks)

Use $Q_1$ to compute the projection of $\mathbf{v}$ onto $R(M)$ where

$$
\mathbf{v} = \left[ \begin{array}{r} 1 \\ 2 \\ 1 \\ 1 \\ -1 \end{array} \right]
$$

Save the result as `proj`.

<!-- grade_id: cell-03f28b32e906b51b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
proj = 
```

<!-- grade_id: cell-4395add049de2804, points: 1, grade: true, solution: false -->
```python
"Verify type of proj. (1 mark)"
assert isinstance(proj,np.ndarray)
print("Problem 3b Test 1: Success!")
```

<!-- grade_id: cell-64e4da933f92fdc5, points: 1, grade: true, solution: false -->
```python
"Verify values of proj. This cell contains hidden tests. (1 mark)"
assert np.allclose(proj[0],0.17647059)
print("Problem 3b Test 2: Success!")
```

### Problem 3c (2 marks)

Find the projection matrix $P$ which projects onto $R(M)^{\perp}$ (the orthogonal complement of the range $R(M)$). Save the result as `P`.

<!-- grade_id: cell-4f28b9e86b3a8f9e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
P = 
```

<!-- grade_id: cell-7e1f979df4e90056, points: 1, grade: true, solution: false -->
```python
"Verify type and size of P. (1 mark)"
assert isinstance(P,np.ndarray) , "P should be a NumPy array."
assert P.shape == (5,5) , "P should be a 5x5 matrix."
print("Problem 3c Test 1: Success!")
```

<!-- grade_id: cell-2b84def6a36feb03, points: 1, grade: true, solution: false -->
```python
"Verify values of P. This cell contains hidden tests. (1 mark)"
assert np.allclose(P[0,0],0.35294118)
print("Problem 3c Test 2: Success!")
```

### Problem 3d (2 marks)

Find the shortest distance from the vector

$$
\mathbf{w} = \left[ \begin{array}{r} 1 \\ 3 \\ 1 \\ 0 \\ 2 \end{array} \right]
$$

to the subspace $R(M)$. Save the result as `distance3`. Use the function [`scipy.linalg.norm`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.norm.html) to compute the norm of a matrix.

<!-- grade_id: cell-5d99588f53797102, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
distance3 = 
```

<!-- grade_id: cell-c656194d0bf57c1c, points: 1, grade: true, solution: false -->
```python
"Verify distance3 is a positive number. (1 mark)"
assert distance3 > 0
print("Problem 3d Test 1: Success!")
```

<!-- grade_id: cell-532b08c4b2a7bcfd, points: 1, grade: true, solution: false -->
```python
"Verify value of distance3. This cell contains hidden tests. (1 mark)"
assert np.round(distance3,1) == 1.7
print("Problem 3d Test 2: Success!")
```