# MATH 210 Assignment 2

### Instructions

* Write your solutions in the cells with `YOUR CODE HERE`.
* You may work with others but submit your own solutions.
* **Do not import** any Python packages such as `math` or `numpy` to complete this assignment.
* Execute the test cells to verify that your solutions pass.
* **This notebook contains hidden tests!** Your solution may not be completely correct even if it passes all tests below.
* Submit this notebook to Canvas.

## Problem 1 (5 marks)

Write a function called `median` which takes a list of numbers `x` and returns the median value (see [Wikipedia:Median](https://en.wikipedia.org/wiki/Median)). If the length of the list `x` is odd then the median is the middle value and if the length is even then the median is the average of the two middle values. For example, the median of `[1,3,-1,6,2]` is `2` and the median of `[1,3,2,6,-1,0]` is `1.5`. Use the built-in function `sorted` to sort the list of values. Use the modulo operator `%` to determine if an integer is odd or even.

This is an exercise in pure Python. Do **not** use any Python packages such as `math` or `numpy` or `scipy` for this exercise.

<!-- grade_id: cell-b8dc729e3aef29ad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-2ac08a19378ec6c1, points: 1, grade: true, solution: false -->
```python
# Test 1: Check median returns the correct datatype. (1 mark)
assert isinstance(median([1,2,3,4]),float)
print("Problem 1 Test 1: Success!")
```

<!-- grade_id: cell-bfb18de96b51b807, points: 1, grade: true, solution: false -->
```python
# Test 2: Check median returns the correct values. (1 mark)
assert abs(median([1.0,2.0,3.0]) - 2.0) < 1e-14
print("Problem 1 Test 2: Success!")
```

<!-- grade_id: cell-5744556726894c69, points: 1, grade: true, solution: false -->
```python
# Test 3: Check median returns the correct values. This cell contains hidden tests. (1 mark)
assert abs(median([1.0,2.0,3.0,4.0]) - 2.5) < 1e-14
print("Problem 1 Test 3: Success!")
```

<!-- grade_id: cell-8c7495ed4cdc39ef, points: 1, grade: true, solution: false -->
```python
# Test 4: Check median returns the correct values. This cell contains hidden tests. (1 mark)
```

<!-- grade_id: cell-0fd8eb31d3170113, points: 1, grade: true, solution: false -->
```python
# Test 5: Check median returns the correct values. This cell contains hidden tests. (1 mark)
```

## Problem 2 (5 marks)

Write a function called `sequence` which takes input parameters `a` and `N` and returns the $N$th value $x_N$ of the recursive sequence

$$
x_0 = 1 \ , \ \ x_{n+1} = \frac{1 + ax_n}{1 + x_n^2}
$$

<!-- grade_id: cell-dfed481424a1bc36, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-305f5c1070b51790, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if sequence returns a float. (1 mark)
assert isinstance(sequence(2,4),float) , "sequence should return a float."
print("Problem 2 Test 1: Success!")
```

<!-- grade_id: cell-fcbd8666766f22fc, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if sequence returns different numbers for different input. (1 mark)
assert abs(sequence(2,5) - sequence(3,10)) > 0
print("Problem 2 Test 2: Success!")
```

<!-- grade_id: cell-c9353f804484e67d, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if sequence returns correct values. (1 mark)
assert abs(sequence(2,10) - 1.323952613955125) < 1e-10
print("Problem 2 Test 3: Success!")
```

<!-- grade_id: cell-97137a53f47984b8, points: 2, grade: true, solution: false -->
```python
# Test 4: Check if sequence returns correct values. This cell contains hidden tests. (2 marks)
```

## Problem 3 (5 marks)

Write a function called `sequence_ab` which takes 3 input parameters `a`, `b` and `N` and returns a Python list `[x0,x1,...,xN]` of length $N+1$ representing the recursive sequence

$$
x_0 = 1 \ , \ \ x_1 = 1 \ , \ \ x_{n+2} = a x_{n+1} + b x_{n}
$$

For example, if $a = b = 1$ and $N = 5$ then the function returns the partial Fibonacci sequence `[1,1,2,3,5,8]`.

<!-- grade_id: cell-00a1f5264fda01d1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-353a2223da490408, points: 1, grade: true, solution: false -->
```python
# Test 1: Check function returns correct datatype. (1 mark)
assert type(sequence_ab(1,1,5)) == list , "Return value should be a list."
print("Problem 3 Test 1: Success!")
```

<!-- grade_id: cell-b1c687cb336730ee, points: 1, grade: true, solution: false -->
```python
# Test 2: Check function returns a list of the correct length. (1 mark)
assert len(sequence_ab(1,-1,23)) == 24 , "Return value should be a list of length 24."
print("Problem 3 Test 2: Success!")
```

<!-- grade_id: cell-867c5eef7bd84ea4, points: 1, grade: true, solution: false -->
```python
# Test 3: Check function returns correct values. (1 mark)"
assert sequence_ab(1,1,5) == [1,1,2,3,5,8] , "Return value should be [1,1,2,3,5,8] when a=b=1 and N=5."
print("Problem 3 Test 3: Success!")
```

<!-- grade_id: cell-4afe32e813ddbcc3, points: 2, grade: true, solution: false -->
```python
# Test 4: Check function returns correct values. This cell contains hidden tests. (2 marks)
```

## Problem 4 (5 marks)

Write a function called `newton_ab` which takes 6 input parameters `a`, `b`, `c`, `x0`, `max_iter` (default value `10`) and `epsilon` (default value `1e-10`) and implements Newton's method with initial guess `x0` and returns an approximation of a solution of the equation $f(x) = 0$ where

$$
f(x) = x^5 + ax^2 + bx + c
$$

The value `max_iter` is the maximum number of iterations of Newton's method to implement. The value `epsilon` is the stopping criteria: $|f(x_n)| < \epsilon$. Use may use the function `newton` defined below.

```python
def newton(f,Df,x0,epsilon,max_iter):
    xn = x0
    for n in range(0,max_iter):
        fxn = f(xn)
        if abs(fxn) < epsilon:
            print('Found solution after',n,'iterations.')
            return xn
        Dfxn = Df(xn)
        if abs(Dfxn) < 1e-12:
            print('Zero derivative. No solution found.')
            return None
        xn = xn - fxn/Dfxn
    print('Exceeded maximum iterations. No solution found.')
    return None
```

<!-- grade_id: cell-23bed5d11ea74a18, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-e153d07912e38c56, points: 1, grade: true, solution: false -->
```python
# Test 1: Check function accepts the right number of input parameters. (1 mark)
assert type(newton_ab(2,-1,-1,1)) == float , "Return value should be a float."
assert type(newton_ab(2,-1,-1,1,10)) == float , "Return value should be a float."
assert type(newton_ab(2,-1,-1,1,10,1e-4)) == float , "Return value should be a float."
print("Problem 4 Test 1: Success!")
```

<!-- grade_id: cell-ead4ec53484a02c8, points: 1, grade: true, solution: false -->
```python
# Test 2: Check function returns the correct datatype when there is a zero derivative. (1 mark)
assert newton_ab(-2,-1,1,1,25,0.001) == None , "Return value should be None when f'(x) = 5x^4 - 4x - 1 and x0 = 1."
print("Problem 4 Test 2: Success!")
```

<!-- grade_id: cell-0bf19de35de2d904, points: 1, grade: true, solution: false -->
```python
# Test 3: Check function returns the correct datatype when exceed maximum iterations. (1 mark)
assert newton_ab(2,-1,-1,10,2,1e-12) == None , "Return value should be None when maximum iterations is exceeded."
print("Problem 4 Test 3: Success!")
```

<!-- grade_id: cell-1821dab43585ca3c, points: 1, grade: true, solution: false -->
```python
# Test 4: Check function returns the correct values. (1 mark)
assert abs(newton_ab(2,-1,-1,1) - 0.8421756416969274) < 1e-12
assert abs(newton_ab(5,4,-3,1,10,1e-12) - 0.4691683039523567) < 1e-12
print("Problem 4 Test 4: Success!")
```

<!-- grade_id: cell-a20d96b980b03f22, points: 1, grade: true, solution: false -->
```python
# Test 5: Check function returns the correct values. This cell contains hidden tests. (1 mark)
```

## Problem 5 (5 marks)

Represent a polynomial $p(x)=c_0+c_1x+c_2x^2+\cdots+c_nx^n$ as a list of coefficients $[c_0,c_1,\dots,c_n]$. Write a function called `poly_critical` which takes 4 input parameters `p`, `a`, `b` and `N` where:

* `p` is a Python list of numbers representing a polynomial $p(x)$
* `a` and `b` are numbers defining an interval $[a,b]$
* `N` is a positive integer

The function `poly_critical` implements $N$ iterations of the bisection method applied to the equation $p'(x)=0$ and returns an approximation of a critical point $x$ where $p'(x)=0$ for $x \in [a,b]$.

For example, if `p = [1,-1,1,0,0,1]` (which represents $p(x) = 1-x+x^2+x^5$), `a=0`, `b=1` and `N=100` then the function returns `0.4212656546685004` which approximates a solution of $5x^4+2x-1=0$.

The function `poly_critical` may terminate in 4 ways:

1. If `len(p) <= 2` (`p` is linear), the function should print `"No critical points."` and return `None`.
2. If either initial endpoint is a critical point, return the endpoint.
3. If values at both endpoints and midpoint have the same sign at any iteration, the function should print `"Bisection method fails."` and return `None`.
4. The function implements N iterations of the bisection method successfully and returns the midpoint of the Nth subinterval.

Use the functions `poly_eval`, `poly_diff` and `bisection` defined below. Note that the function `bisection` satisfies the condition 2,3,4 above for appropriate input.

```python
def poly_eval(p,a):
    return sum([p[n]*a**n for n in range(0,len(p))])

def poly_diff(p):
    if len(p) > 1:
        dpdx = [p[n]*n for n in range(1,len(p))]
    else:
        dpdx = [0]
    return dpdx
    
def bisection(f,a,b,N):
    if f(a) == 0:
        print("Found exact solution at endpoint x=a.")
        return a
    if f(b) == 0:
        print("Found exact solution at endpoint x=b.")
        return b
    if f(a)*f(b) >= 0:
        print("Bisection method fails.")
        return None
    an = a
    bn = b
    for n in range(1,N+1):
        mn = (an + bn)/2
        fmn = f(mn)
        if f(an)*fmn < 0:
            bn = mn
        elif f(bn)*fmn < 0:
            an = mn
        elif fmn == 0:
            print("Found exact solution.")
            return mn
        else:
            print("Bisection method fails.")
            return None
    return (an + bn)/2
```

<!-- grade_id: cell-1d2fdd2df0746fdf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a7c018588f6973a2, points: 1, grade: true, solution: false -->
```python
# Test 1: Check function returns the correct value when p(x) is linear. (1 mark)
assert poly_critical([1,1],0,1,100) == None , "Return value should be a None if p(x) is linear."
print("Problem 5 Test 1: Success!")
```

<!-- grade_id: cell-144fdd5ea907bd8f, points: 2, grade: true, solution: false -->
```python
# Test 2: Check function returns the correct value when critical points at the endpoints. (2 marks)
assert abs(poly_critical([0,-3,0,1],0,1,100) - 1) < 1e-12 , "Return values should be 1.0 when p(x) = x^3 - 3x and a=0, b=1."
assert abs(poly_critical([0,-3,0,1],-1,0,100) - -1) < 1e-12 , "Return values should be -1.0 when p(x) = x^3 - 3x and a=-1, b=0."
print("Problem 5 Test 2: Success!")
```

<!-- grade_id: cell-44f11f0f6b978db4, points: 2, grade: true, solution: false -->
```python
# Test 3: Check function returns the correct datatype and value. This cell contains hidden tests. (2 marks)
```

## Problem 6 (5 marks)

Given a finite sequence of number $[a_0,a_1,\dots,a_n]$ (of length $n+1$), define a new finite sequence $[b_0,b_1,\dots, b_n]$ by the recursive formula

$$
b_n = a_{n} + \frac{1}{b_{n-1}}
$$

For example:

\begin{align*}
b_0 &= a_0 \\
b_1 &= a_{1} + \frac{1}{b_0} = a_{1} + \frac{1}{a_0} \\
b_2 &= a_{2} + \frac{1}{b_1} = a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}\\
b_3 &= a_{3} + \frac{1}{b_2} = a_3 + \frac{1}{a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}} \\
& \vdots \\
b_n &= a_{n} + \frac{1}{b_{n-1}} = a_n + \frac{1}{\ddots + \frac{1}{a_0}} 
\end{align*}

Write a function called `sequence_to_fraction` which takes one input parameter `a` (a Python list of positive integers $[a_0,a_1,\dots,a_n]$) and returns the last number $b_n$ in the sequence defined above

$$
b_n = a_n + \frac{1}{a_{n-1} + \frac{1}{\ddots + \frac{1}{a_0}}}
$$




<!-- grade_id: cell-9c6015ad6263af5a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-56b9cd8cb0a27347, points: 1, grade: true, solution: false -->
```python
# Test 1: Check function returns the correct datatype. (1 mark)
assert type(sequence_to_fraction([1,1,1,1,1,1,1,1,1,1,1])) == float , "Return value should be a float."
print("Problem 6 Test 1: Success!")
```

<!-- grade_id: cell-6f7304bcbb8d2aa2, points: 1, grade: true, solution: false -->
```python
# Test 2: Check function returns the correct values. (1 mark)
assert abs(sequence_to_fraction([1,1,1]) - 1.5) < 1e-12
print("Problem 6 Test 2: Success!")
```

<!-- grade_id: cell-3791ac9e593daa46, points: 1, grade: true, solution: false -->
```python
# Test 3: Check function returns the correct values. (1 mark)
assert abs(sequence_to_fraction([1,1,1,1,1,1,1,1,1,1,1]) - 1.6179775280898876) < 1e-12
print("Problem 6 Test 3: Success!")
```

<!-- grade_id: cell-a7cdbfb491c94b6d, points: 2, grade: true, solution: false -->
```python
# Test 4: Check function returns the correct values. This cell contains hidden tests. (2 marks)
```