# MATH 210 Assignment 1

### Instructions

* Write your solutions in the cells with `YOUR CODE HERE`.
* You may work with others but submit your own solutions.
* **Do not import** any Python packages such as `math` or `numpy` to complete this assignment.
* Execute the test cells to verify that your solutions pass.
* **This notebook contains hidden tests!** Your solution may not be completely correct even if it passes all tests below.
* Submit this notebook to Canvas.

### Problem 1 (4 marks)

Consider the infinite series formula

$$
\sum_{k = 1}^{\infty} \frac{(-1)^{k+1}}{(2k - 1)^5} = \frac{5 \pi^5}{1536}
$$

We can rearrange the formula as

$$
\pi = \left(\frac{1536}{5} \sum_{k = 1}^{\infty} \frac{(-1)^{k+1}}{(2k - 1)^5} \right)^{1/5}
$$

Use a list comprehension and the builtin function `sum` to compute the first 10 terms of the series

$$
\left(\frac{1536}{5} \sum_{k = 1}^{10} \frac{(-1)^{k+1}}{(2k - 1)^5} \right)^{1/5}
$$

and assign the value to the variable `P1`. (Note: the result `P1` is an approximation of $\pi$.)

<!-- grade_id: cell-60575b628ba6c404, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-58beebcb9186fd19, points: 1, grade: true, solution: false -->
```python
"Check if P1 is a float. (1 mark)"
assert isinstance(P1,float) , "P1 should be a float."
print("Problem 1 Test 1: Success!")
```

<!-- grade_id: cell-9823d1c68053d38d, points: 1, grade: true, solution: false -->
```python
"Check if P1 is an approximation of pi. (1 mark)"
assert abs(P1 - 3.14159) < 0.00001 , "P1 should be an approximation of pi."
print("Problem 1 Test 2: Success!")
```

<!-- grade_id: cell-97e2685b56db8621, points: 2, grade: true, solution: false -->
```python
"Check if P1 has the correct value. This cell contains hidden tests. (2 marks)"
assert abs(P1 - 3.14159) < 0.00001 , "P1 should be an approximation of pi."
print("Problem 1 Test 3: Success!")
```

### Problem 2 (3 marks)

Write a list comprehension to create a list with values

$$
[-101,-97,-93,-89,\dots,87,91,95,99]
$$

Save the result as `P2`.

<!-- grade_id: cell-6f38c5c2d9810a0d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b1156af4c3187285, points: 1, grade: true, solution: false -->
```python
"Check if P2 is a list of length 51. (1 mark)"
assert isinstance(P2,list) , "P2 should be a list."
assert len(P2) == 51, "P2 should be a list of length 51."
print("Problem 2 Test 1: Success!")
```

<!-- grade_id: cell-6c2c0365b00a66e2, points: 2, grade: true, solution: false -->
```python
"Check if P2 has the correct values in the list. This cell contains hidden tests. (2 marks)"
assert P2[:4] == [-101,-97,-93,-89], "First 4 values of P2 are -101, -97, -93, -89."
print("Problem 2 Test 2: Success!")
```

### Problem 3 (4 marks)

Represent a vector $\mathbf{v} \in \mathbb{R}^3$ as a list of numbers `[a,b,c]`. The [cross product](https://en.wikipedia.org/wiki/Cross_product) of vectors is given by the formula

$$
[ a , b , c ] \times [ x , y , z ] = [ bz - cy , cx - az , ay - bx ]
$$

Write a function called `cross` which takes 2 input parameters `u` and `v` (lists of length 3 representing vectors $\mathbf{u}$ and $\mathbf{v}$), and returns the cross product $\mathbf{u} \times \mathbf{v}$ as a list of length 3.

<!-- grade_id: cell-c98f723b81172d79, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-71b70efd48c8a261, points: 1, grade: true, solution: false -->
```python
"Check if cross accepts 2 input parameters. (1 mark)"
cross([1,2,3],[4,5,6])
assert True , "cross does not run."
print("Problem 3 Test 1: Success!")
```

<!-- grade_id: cell-08e4ea2222319e00, points: 1, grade: true, solution: false -->
```python
"Check if cross returns a list. (1 mark)"
w = cross([1,2,3],[4,5,6])
assert isinstance(w,list) , "cross should return a list."
print("Problem 3 Test 2: Success!")
```

<!-- grade_id: cell-af4dcfcca94d72ce, points: 2, grade: true, solution: false -->
```python
"Check if cross returns correct values. This cell contains hidden tests. (2 marks)"
u = [1,0,0]
v = [0,1,0]
w = cross(u,v)
assert w == [0,0,1] , "cross([1,0,0],[0,1,0]) should return [0,0,1]."
print("Problem 3 Test 3: Success!")
```

### Problem 4 (4 marks)

Write a function called `seriesN` which takes 2 input parameters `x` and `N` and returns the sum

$$
\sum_{n=0}^{N} \frac{x^2}{2^n + 1}
$$

<!-- grade_id: cell-2c498a7bf0750d90, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-6cfee83dfdd14fe7, points: 1, grade: true, solution: false -->
```python
"Check if seriesN accepts 2 input parameters. (1 mark)"
seriesN(0.2,10)
assert True , "seriesN does not run."
print("Problem 4 Test 1: Success!")
```

<!-- grade_id: cell-33ffba19403e82f6, points: 1, grade: true, solution: false -->
```python
"Check if seriesN returns a float. (1 mark)"
result = seriesN(0.2,10)
assert isinstance(result,float) , "seriesN should return a float."
print("Problem 4 Test 2: Success!")
```

<!-- grade_id: cell-42899dae299322f8, points: 2, grade: true, solution: false -->
```python
"Check if seriesN returns correct values. This cell contains hidden tests. (2 marks)"
assert abs(seriesN(1,1) - 5/6) < 1e-12, "seriesN(1,1) should return 5/6."
print("Problem 4 Test 3: Success!")
```

### Problem 5 (5 marks)

Represent a polynomial $p(x) = a_0 + a_1 x + \cdots + a_d x^d$ as a list of numbers `[a0,a1,...,ad]`. Write a function called `poly_integral` which takes a list `p` representing a polynomial $p(x) = a_0 + a_1 x + \cdots + a_d x^d$, and numbers `a` and `b`, and returns the value of definite integral

$$
\int_a^b p(x) \, dx
$$

<!-- grade_id: cell-10709266a0c473a7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-112c6d82a168f976, points: 1, grade: true, solution: false -->
```python
"Check if poly_diff accepts 3 input parameters. (1 mark)"
poly_integral([1,2,3],0,1)
assert True , "poly_integral does not run."
print("Problem 5 Test 1: Success!")
```

<!-- grade_id: cell-661f4358ce446496, points: 1, grade: true, solution: false -->
```python
"Check if poly_integral returns a float when p = 1+x and a=0,b=1. (1 mark)"
result = poly_integral([1,1],0,1)
assert isinstance(result,float) , "poly_integral should return a float."
print("Problem 5 Test 2: Success!")
```

<!-- grade_id: cell-f84d04b6ff99ab1d, points: 3, grade: true, solution: false -->
```python
"Check if poly_integral returns the correct values. This cell contains hidden tests. (3 marks)"
result = poly_integral([1,1],0,1)
assert abs(result - 1.5) < 1e-10 , "poly_integral([1,1],0,1) should return 1.5."
print("Problem 5 Test 3: Success!")
```