# MATH 210 Assignment 3

### Instructions

* Write your solutions in the cells with `YOUR CODE HERE`.
* You may work with others but submit your own solutions.
* Execute the test cells to verify that your solutions pass.
* **This notebook contains hidden tests!** Your solution may not be completely correct even if it passes all tests below.
* Submit this notebook to Canvas.

```python
import numpy as np
import matplotlib.pyplot as plt
```

## Problem 1 (6 marks)

Write a function called `fourier` which takes input parameters `c` and `num` where `c` is a Python list of numbers `c = [c0,c1,...,cN]` and `num` is a postive integer. The function `fourier` plots the function

$$
f(x) = \frac{c_0}{2} + \sum_{k=1}^N c_k \cos(2 \pi k x) 
$$

over the interval $[0,1]$ using `num` equally spaced points from `0` to `1`. (In other words, use `np.linspace(0,1,num)` for the `x` values.) The function also returns the `x` and `y` values of the plot in a matrix with 2 columns with `x` values in the first column and `y` values in the second. The function `np.column_stack` may be helpful. (Note that the function is a (partial) [Fourier cosine series](https://en.wikipedia.org/wiki/Fourier_sine_and_cosine_series).)

For example:

```
fourier([1,2,4,8,16],100)
```

creates the figure

<img src="http://www.math.ubc.ca/~pwalls/data/a3prob1.png" style="width: 500px">

<!-- grade_id: cell-9d61d16703ed95e0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-9f9691b3351f1558, points: 1, grade: true, solution: false -->
```python
"Check fourier returns the correct datatype. (1 mark)"
assert isinstance(fourier([1,1,5,1,5],200),np.ndarray) , "Return value should be a NumPy array."
print("Problem 1 Test 1: Success!")
```

<!-- grade_id: cell-7c3743e922404df7, points: 2, grade: true, solution: false -->
```python
"Check fourier returns the correct size of NumPy array. (2 marks)"
assert fourier([1,2,3],200).shape == (200,2) , "Return value should be matrix with size num by 2."
print("Problem 1 Test 2: Success!")
```

<!-- grade_id: cell-b5887c9275dd0c44, points: 3, grade: true, solution: false -->
```python
"Check fourier returns the correct values. This cell contains hidden tests. (3 marks)"
epsilon = 1e-8
M = fourier([1,-1,2,-1,1],200)
ans1 = np.array([[0., 1.5],[0.00502513, 1.49303019],[0.01005025, 1.47222285]])
ans2 = np.array([[0.50251256, 5.49576411],[0.50753769, 5.46194589],[0.51256281, 5.39467572]])
ans3 = np.array([[0.98994975, 1.47222285],[0.99497487, 1.49303019],[1., 1.5]])
assert np.max(np.abs(M[:3,:] - ans1)) < epsilon
assert np.max(np.abs(M[100:103,:] - ans2)) < epsilon
assert np.max(np.abs(M[-3:,:] - ans3)) < epsilon
print("Problem 1 Test 3: Success!")
```

## Problem 2 (4 marks)

Write a function called `k_product` which takes input parameters `x` and `N` and returns the product

$$
x\prod_{k=1}^N \left( 1 - \frac{x^2}{k^2 \pi^2} \right)
$$

Use the function `np.prod` with NumPy arrays. Note the infinite product satisfies

$$
\sin(x) = x \prod_{k=1}^{\infty} \left( 1 - \frac{x^2}{k^2 \pi^2} \right)
$$

Do **not** use a `for` loop. Use NumPy array functions and operations.

<!-- grade_id: cell-c5499ece79c20e61, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-28f775de1774745d, points: 1, grade: true, solution: false -->
```python
"Check that k_product returns the correct datatype. (1 mark)"
assert isinstance(k_product(1,2),np.float64) , "Return value should be a NumPy float."
print("Problem 2 Test 1: Success!")
```

<!-- grade_id: cell-94f5e49183b4fecc, points: 1, grade: true, solution: false -->
```python
"Check that k_product returns the correct values. (1 mark)"
epsilon = 1e-10
assert np.abs(k_product(1,1) - (1 - 1/np.pi**2)) < epsilon
print("Problem 2 Test 2: Success!")
```

<!-- grade_id: cell-b8a6d21cfd4852cc, points: 2, grade: true, solution: false -->
```python
"Check that k_product returns the correct values. This cell contains hidden tests. (2 marks)"
epsilon = 1e-10
assert np.abs(k_product(np.pi/4,50) - 0.7079824300576587) < epsilon
print("Problem 2 Test 3: Success!")
```

## Problem 3 (5 marks)

Write a function called `mn_integral` which takes input parameters `a`, `b`, `m`, `n` and `N` and returns the Riemann sum approximation (with midpoints and `N` subintervals of equal length) of the integral
$$
\int_a^b x^m \cos(nx) \, dx
$$

<!-- grade_id: cell-0cf55684009b90e1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b7675f8540aa93b6, points: 1, grade: true, solution: false -->
```python
"Check that mn_integral returns the correct type. (1 mark)"
assert isinstance(mn_integral(0,2,1,1,10),np.float64) , "Return type should be a NumPy float."
print("Problem 3 Test 1: Success!")
```

<!-- grade_id: cell-97d7cd39e5d0c1ea, points: 2, grade: true, solution: false -->
```python
"Check that mn_integral returns the correct values. (2 marks)"
epsilon = 1e-8
assert np.abs(mn_integral(0,np.pi/4,1,1,1) - 0.28494767187968045) < epsilon
print("Problem 3 Test 2: Success!")
```

<!-- grade_id: cell-602410d655a8bbca, points: 2, grade: true, solution: false -->
```python
"Check that mn_integral returns the correct values. (2 marks)"
assert mn_integral(0,2,1,0,1) == 2
print("Problem 3 Test 3: Success!")
```

##  Problem 4 (5 marks)

Write a function called `erf` which takes input parameters `x` and `error` and returns the approximation of the integral

$$
\int_0^x e^{-t^2} dt
$$

given by the trapezoid method with $N$ subintervals of equal length which guarantees the error is less than `error`. For $K_2$, use the exact value of the maximum of $|f''(x)|$ for $x \in [0,\infty)$.

<!-- grade_id: cell-b170f87ca3d3aec0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-a885827d438220b3, points: 1, grade: true, solution: false -->
```python
"Check that erf returns the correct type. (1 mark)"
assert isinstance(erf(2,0.001),np.float64) , "Return value should be a NumPy float."
print("Problem 4 Test 1: Success!")
```

<!-- grade_id: cell-4821e7265743adf9, points: 2, grade: true, solution: false -->
```python
"Check that erf returns the correct values. (2 marks)"
epsilon = 1e-10
assert np.abs(erf(1,0.01) - 0.744368339763667) < epsilon
print("Problem 4 Test 2: Success!")
```

<!-- grade_id: cell-5188eda2ac651f00, points: 2, grade: true, solution: false -->
```python
"Check that erf returns the correct values. This cell contains hidden tests. (2 marks)"
epsilon = 1e-10
assert np.abs(erf(1,0.0001) - 0.7467876578237478) < epsilon
print("Problem 4 Test 3: Success!")
```

##  Problem 5 (5 marks)

Write a function called `Ei` which takes input parameters `x` and `error` and returns the approximation of the integral

$$
\int_1^x \frac{e^{-t}}{t} dt
$$

given by the trapezoid method with $N$ subintervals of equal length which guarantees the error is less than `error`. For $K_2$, use the exact value of the maximum of $|f''(x)|$ for $x \in [1,\infty)$.

<!-- grade_id: cell-089c21c078558af6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-0286fc9d045097fc, points: 1, grade: true, solution: false -->
```python
"Check that Ei returns the correct type. (1 mark)"
assert isinstance(Ei(2,0.001),np.float64) , "Return value should be a NumPy float."
print("Problem 5 Test 1: Success!")
```

<!-- grade_id: cell-50cbb1d431ff52db, points: 1, grade: true, solution: false -->
```python
"Check that Ei returns the correct values. (1 mark)"
assert Ei(1,0.01) == 0
print("Problem 5 Test 2: Success!")
```

<!-- grade_id: cell-713488b8daf3f5e7, points: 3, grade: true, solution: false -->
```python
"Check that Ei returns the correct values. This cell contains hidden tests. (3 marks)"
epsilon = 1e-10
assert np.abs(Ei(5,0.001) - 0.21833350347126337) < epsilon
print("Problem 5 Test 3: Success!")
```