# MATH 257/316 Python Assignment 1

* See [Python for UBC Math](https://ubcmath.github.io/python/) for an introduction to Python and Jupyter
* See [MATH 316 Jupyter Notebooks](https://ubcmath.github.io/MATH316/) to see more examples
* Run the tests to verify your solutions
* There are **hidden tests** therefore your solutions may not be entirely correct even if they pass the tests below
* Submit your `.ipynb` notebook file to Canvas (download from Syzygy to your machine and upload to Canvas)

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
```

## Problem 1 (5 marks)

Consider the equation

$$
u_t = \alpha^2 u_{xx} \ , \ \ 0 \leq x \leq \pi \ , \ \ t \geq 0
$$

with $\alpha^2 = 2$, initial condition $u(x,0) = \cos(x/2)$ and boundary conditions $u_x(0,t) = u_x(\pi,t)=0$. The Fourier series solution is of the form

$$
u(x,t) = \frac{a_0}{2} + \sum_{n=1}^{\infty} a_n \cos(nx) e^{-2 n^2 t}
$$

Find a formula for the coefficients $a_n$. Enter the formula as a Python function called `a` with input parameter `n`.

<!-- grade_id: cell-a1b4f2d27e9c07d9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
a = lambda n: # enter formula for a(n)
```

<!-- grade_id: cell-b2e0934d3a596790, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify that a is defined as a function. (1 mark)
assert callable(a) == True ,  "a should be a Python function."
print("Test 1: Success!")
```

<!-- grade_id: cell-f39841d13f1e3830, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify that a(0), a(1) and a(3) are positive numbers. (1 mark)
assert a(0) > 0 ,  "a(0) should be positive."
assert a(1) > 0 ,  "a(1) should be positive."
assert a(3) > 0 ,  "a(3) should be positive."
print("Test 2: Success!")
```

<!-- grade_id: cell-2d0243d0f956d23c, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify a(0) and a(4) are correct values. (1 mark)
assert abs(a(0) - 4/np.pi) < 1e-12 ,  "a(0) should be 4/pi."
assert abs(a(4) + 0.020210151503732742) < 1e-12 ,  "a(4) should be -0.020210151503732742."
print("Test 3: Success!")
```

<!-- grade_id: cell-d4620c8bc71abbdf, points: 2, grade: true, solution: false -->
```python
# Test 4: Verify function a(n) returns correct values. This cell contains hidden tests. (2 marks)
```

The following code plots the partial sum

$$
u(x,t) \approx \frac{a_0}{2} + \sum_{n=1}^{N} a_n \cos(nx) e^{-2 n^2 t}
$$

up to $N=100$ for each $t=0.0,0.1,0.2,\dots,2.0$:

```python
N = 100
x = np.linspace(0,np.pi,200)
for t in np.linspace(0,2,21):
    u = a(0)/2 + sum([a(n)*np.cos(n*x)*np.exp(-2*n**2*t) for n in range(1,N+1)])
    plt.plot(x,u,'C0')
plt.grid(True)
plt.show()
```

The plot at time $t=0$ should match the initial condition $u(x,0) = \cos(x/2)$. The plot at time $t=2$ should be close to the steady state solution.

## Problem 2 (5 marks)

The function `heatFTCSN` computes the approximation of the heat equation with Neumann boundary conditions using the foward-time-central-space finite difference method. Compute the numerical solution for the equation in Problem 1 above using $t_f=2$, $N=20$ and $M=5000$. Save the numerical approximation as a matrix `U`

```python
def heatFTCSN(alpha,L,f,Q0,QL,tf,N,M):
    dx = L/N
    dt = tf/M
    x = np.linspace(0,L,N+1)
    U = np.zeros((N+1,M+1))
    U[:,0] = f(x)
    r = alpha*dt/dx**2
    A = np.zeros((N+1,N+1))
    A[0,0] = 1 - 2*r
    A[0,1] = 2*r
    A[N,N] = 1 - 2*r
    A[N,N-1] = 2*r
    for n in range(1,N):
        A[n,n-1] = r
        A[n,n] = 1 - 2*r
        A[n,n+1] = r
    q = np.zeros(N+1)
    q[0] = -2*alpha**2*dt/dx*Q0
    q[N] = 2*alpha**2*dt/dx*QL
    for k in range(M):
        U[:,k+1] = A@U[:,k] + q
    return U
```

<!-- grade_id: cell-23f9eb94e5d35439, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
U = 
```

<!-- grade_id: cell-fe9ab9fd69999927, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify U is a NumPy array of size 21 x 5001. (1 mark)
assert isinstance(U,np.ndarray) == True ,  "U should be a NumPy array."
assert U.shape == (21,5001),  "U should a matrix of size 21 by 5001."
print("Test 1: Success!")
```

<!-- grade_id: cell-c27253a2649d8aa0, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify initial condition is satisfied: u(x,0) = cos(x/2). (1 mark)
x = np.linspace(0,np.pi,21)
fx = np.cos(x/2)
assert np.allclose(U[:,0],fx),  "u(x,0) = cos(x/2)."
print("Test 2: Success!")
```

<!-- grade_id: cell-42377cab88b29720, points: 1, grade: true, solution: false -->
```python
# Test 3: Verify values at t = 2. (1 mark)
u2 = np.array([0.66154164, 0.66123082, 0.66030602, 0.65878999, 0.65672005,
                0.65414713, 0.65113458, 0.64775654, 0.64409617, 0.64024358,
                0.63629362, 0.63234355, 0.62849064, 0.62482977, 0.6214511 ,
                0.61843784, 0.61586423, 0.61379366, 0.61227713, 0.61135201,0.61104108])
assert np.allclose(U[:,-1],u2) ,  "u(0,2) = ."
print("Test 3: Success!")
```

<!-- grade_id: cell-d02677ded3467f70, points: 2, grade: true, solution: false -->
```python
# Test 4: Verify values of U are correct. This cell contains hidden tests. (2 marks)
```

The following code plots the finite difference solution $u_{n,k}$ for $k=0,250,500,750,\dots,5000$.

```python
x = np.linspace(0,np.pi,21)
dt = 2/5000
for t in np.linspace(0,2,21):
    k = int(t/dt)
    u = U[:,k]
    plt.plot(x,u,'C0')
plt.grid(True)
plt.show()
```

The plot at time $t=0$ should match the initial condition $u(x,0) = \cos(x/2)$. The plot at time $t=2$ should be close to the steady state solution. The figure should match the figure in Problem 1.