# MATH 257/316 Python Assignment 2

* See [Python for UBC Math](https://ubcmath.github.io/python/) for an introduction to Python and Jupyter
* See [MATH 316 Jupyter Notebooks](https://ubcmath.github.io/MATH316/) to see more examples
* Run the tests to verify your solutions
* There are **hidden tests** therefore your solutions may not be entirely correct even if they pass the tests below
* Submit your `.ipynb` notebook file to Canvas (download from Syzygy to your machine and upload to Canvas)

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
import scipy.integrate as spi
```

## Piecewise Functions with NumPy

The function `np.piecewise` creates a piecewise defined funciton. See the [documentation](https://numpy.org/doc/stable/reference/generated/numpy.piecewise.html). Let's use `np.piecewise` to create the function step function:

$$
u_{a,b}(x) = \left\{ \begin{array}{cc} 0 & x <  a \\ 1 & a \leq x \leq b \\ 0 & x > b \end{array} \right.
$$

```python
u = lambda x,a,b: np.piecewise(x,[x < a,x >= a,x > b],[0,1,0])
```

Plot the function $u_{a,b}(x)$ for $a=-1,b=1$ on the interval $[-2,2]$.

```python
x = np.linspace(-2,2,200)
y = u(x,-1,1)
plt.plot(x,y)
plt.show()
```

Use the step function $u_{a,b}(x)$ to construct the even extension $f_{even}(x)$ to the interval $[-2,2]$ of the following function:

$$
f(x) = \left\{ \begin{array}{cc} 0 & 0 \leq x < 1 \\ x & 1 \leq x \leq 2 \end{array} \right.
$$

```python
feven = lambda x: -x*u(x,-2,-1) + x*u(x,1,2)
```

Plot the function $f_{even}(x)$ on the interval $[-2,2]$:

```python
x = np.linspace(-2,2,200)
y = feven(x)
plt.plot(x,y)
plt.grid(True), plt.axis('equal')
plt.show()
```

## Problem 1 (4 marks)

Use the step function $u_{a,b}(x)$ to construct the **odd** extension $f_{odd}(x)$ of the following function to the interval $[-2,2]$:

$$
f(x) = \left\{ \begin{array}{cc} 0 & 0 \leq x < 1 \\ 1 - \cos(2 \pi x) & 1 \leq x \leq 2 \end{array} \right.
$$

Save $f_{odd}(x)$ as a Python function called `fodd`.

<!-- grade_id: cell-980a4992f4bee72e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
fodd = lambda x: 
```

<!-- grade_id: cell-f9d6db0c258c9cff, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify that fodd is defined as a function. (1 mark)
assert callable(fodd) == True ,  "fodd should be a Python function."
print("Test 1: Success!")
```

<!-- grade_id: cell-970c83ad7c4766e3, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify values of fodd. (1 mark)
assert fodd(-1.5) == -2 ,  "food(-1.5) should be -2."
assert fodd(-0.5) == 0 ,  "food(-0.5) should be 0."
assert fodd(0.5) == 0 ,  "food(0.5) should be 0."
assert fodd(1.5) == 2 ,  "food(-1.5) should be 2."
print("Test 2: Success!")
```

<!-- grade_id: cell-659fb0e585bc5442, points: 2, grade: true, solution: false -->
```python
# Test 3: Verify function fodd returns correct values. This cell contains hidden tests. (2 marks)
```

Plot the function $f_{odd}(x)$ over the interval $[-2,2]$:

```python
x = np.linspace(-2,2,200)
plt.plot(x,fodd(x))
plt.show()
```

## Problem 2 (4 marks)

The function `fourier` below computes the Fourier coefficients for a function $f(x)$ on the interval $[-L,L]$. See [MATH 316 > Fourier Series](https://ubcmath.github.io/MATH316/fourier.html).

```python
def fourier(f,L,N):
    a = np.zeros(N)
    b = np.zeros(N)
    I,_ = spi.quad(f,-L,L)
    a0 = 1/L*I
    for n in range(1,N+1):
        integrand = lambda x: f(x)*np.cos(n*np.pi*x/L)
        I,_ = spi.quad(integrand,-L,L)
        a[n-1] = 1/L*I
        integrand = lambda x: f(x)*np.sin(n*np.pi*x/L)
        I,_ = spi.quad(integrand,-L,L)
        b[n-1] = 1/L*I
    return a0,a,b
```

Use the function `fourier` to compute the Fourier coefficients of the **odd** extension $f_{odd}(x)$ to the interval $[-2,2]$ of the function $f(x)$ defined in Problem 1 above. Since $f_{odd}(x)$ is an odd funciton we have $a_n = 0$ for all $n$. Save the coefficients $b_1,b_2,b_3,b_4,b_5$ as `b1`, `b2`, `b3`, `b4`, `b5` respectively.

<!-- grade_id: cell-ef61519740e0b38a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
b1 = 
b2 = 
b3 = 
b4 = 
b5 = 
```

<!-- grade_id: cell-658ad66cfb3f6420, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify that b1,b2,b3,b4,b5 are numbers. (1 mark)
assert isinstance(b1,float) ,  "b1 should be a floating point number."
assert isinstance(b2,float) ,  "b2 should be a floating point number."
assert isinstance(b3,float) ,  "b3 should be a floating point number."
assert isinstance(b4,float) ,  "b4 should be a floating point number."
assert isinstance(b5,float) ,  "b5 should be a floating point number."
print("Test 1: Success!")
```

<!-- grade_id: cell-e3214a23de410330, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify first 2 decimal places of b1,b2,b3,b4,b5. (1 mark)
assert np.round(b1,2) == .68 ,  "b1 rounded to 2 decimal places should be 0.68."
assert np.round(b2,2) == -.85 ,  "b1 rounded to 2 decimal places should be -0.85."
assert np.round(b3,2) == .49 ,  "b1 rounded to 2 decimal places should be 0.49."
assert np.round(b4,2) == .00 ,  "b1 rounded to 2 decimal places should be 0.00."
assert np.round(b5,2) == -.23 ,  "b1 rounded to 2 decimal places should be -0.23."
print("Test 2: Success!")
```

<!-- grade_id: cell-8c4836c120137004, points: 2, grade: true, solution: false -->
```python
# Test 3: Verify b1,b2,b3,b4,b5 are correct values. This cell contains hidden tests. (2 marks)
```

Plot the partial sine series of the odd extension:

```python
x = np.linspace(-2,2,200)
bs = [b1,b2,b3,b4,b5]
y = sum([bs[n-1]*np.sin(np.pi*n*x/2) for n in range(1,6)])
plt.plot(x,fodd(x),x,y)
plt.show()
```

## Problem 3 (6 marks)

Consider the equation

$$
u_t = u_{xx} \ , \ \ 0 \leq x \leq 1 \ , \ \ t \geq 0
$$

with mixed boundary conditions $u(0,t) = u_x(1,t)=0$, and initial condition $u(x,0) = f(x)$ where 

$$
f(x) = \left\{ \begin{array}{cc} 0 & 0 \leq x < 0.4 \\ 1 & 0.4 \leq x \leq 0.8 \\ 0 & 0.8 < x \leq 1 \end{array} \right.
$$

The function `heat` below implements the forward-time-central-space (FTCS) finite difference method for the heat equation with constant Dirichlet/Neumann boundary conditions. See [MATH 316 Jupyter Notebooks > Heat Equation](https://ubcmath.github.io/MATH316/heat.html#ftcs-with-mixed-bcs).

```python
def heat(alpha,L,f,BCtype0,BC0,BCtypeL,BCL,tf,N,M):
    
    dx = L/N
    dt = tf/M
    r = alpha*dt/dx**2
    
    x = np.linspace(0,L,N+1)
    U = np.zeros((N+1,M+1))
    U[:,0] = f(x)
    A = np.zeros((N+1,N+1))
    q = np.zeros(N+1)
    
    if (BCtype0 not in ['D','N']) or (BCtypeL not in ['D','N']):
        raise Exception("Expecting boundary conditions of type 'D' or 'N'.")
    if BCtype0 == 'D':
        U[0,0] = BC0
        A[0,0] = 1
    if BCtype0 == 'N':
        A[0,0] = 1 - 2*r
        A[0,1] = 2*r
        q[0] = -2*alpha**2*dt/dx*BC0
    if BCtypeL == 'D':
        U[N,0] = BCL
        A[N,N] = 1
    if BCtypeL == 'N':
        A[N,N] = 1 - 2*r
        A[N,N-1] = 2*r
        q[N] = 2*alpha**2*dt/dx*BCL
    
    for n in range(1,N):
        A[n,n-1] = r
        A[n,n] = 1 - 2*r
        A[n,n+1] = r
        
    for k in range(M):
        U[:,k+1] = A@U[:,k] + q
        
    return U
```

**Part a.** Use the function `heat` to compute the numerical solution up to $t_f = 0.25$ for $N=50$ and $M=2000$. Save the result as `Ua`.

<!-- grade_id: cell-693c6e25ee8089f9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
Ua = 
```

<!-- grade_id: cell-3112a8ecf96f48e3, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify Ua is a NumPy array of size 51 x 2001. (1 mark)
assert isinstance(Ua,np.ndarray) ,  "Ua should be a NumPy array."
assert Ua.shape == (51,2001) ,  "Ua should be a NumPy array of size 51 by 2001."
print("Test 1: Success!")
```

<!-- grade_id: cell-6fa7d02f010f93dd, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify first column of Ua corresponds to f(x). (1 mark)
assert Ua[0,0] == 0 ,  "Ua[0,0] should be f(0) = 0."
assert Ua[30,0] == 1 ,  "Ua[30,0] should be f(0.6) = 1."
assert Ua[45,0] == 0 ,  "Ua[30,0] should be f(0.9) = 0."
print("Test 2: Success!")
```

<!-- grade_id: cell-ad1554bf123ee131, points: 2, grade: true, solution: false -->
```python
# Test 3: Verify Ua has the correct values. This cell contains hidden tests. (2 marks)
```

Plot the numerical solution:

```python
plt.imshow(Ua,aspect='auto',cmap='jet'), plt.colorbar()
plt.show()
```

We know that the steady state is $u(x,t) \to 0$ as $t \to \infty$. Does this agree with the plot above?

Plot the solution for $t=0.00,0.01,0.02,0.03,0.04,0.05$:

```python
nsteps = 5
tstep = 80
x = np.linspace(0,1,51)
for k in range(nsteps+1):
    t = k*0.01
    plt.plot(x,Ua[:,k*tstep],label='t={:.2f}'.format(t))
plt.legend(loc='upper left')
plt.grid(True)
plt.show()
```

**Part b.** Use the function `heat` to compute the numerical solution up to $t_f = 0.25$ for $N=50$ and and find the smallest $M$ such that the numerical solution is stable (see [numerical stability of FTCS scheme](https://en.wikipedia.org/wiki/FTCS_scheme#Stability).). The value $M$ corresponds to a time step $\Delta t$. Save the value $M$ as `M` and $\Delta t$ as `dt`.

<!-- grade_id: cell-4e61d917431bf34a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
M = 
dt = 
```

<!-- grade_id: cell-224ef79daae534c9, points: 1, grade: true, solution: false -->
```python
# Test 1: Verify M and dt are positive numbers. (1 mark)
assert M > 0 ,  "M should be a positive number."
assert dt > 0 ,  "dta should be a positive number."
print("Test 1: Success!")
```

<!-- grade_id: cell-8d01d26fcc72a38c, points: 1, grade: true, solution: false -->
```python
# Test 2: Verify M and dt are the correct values. This cell contains hidden tests. (2 marks)
```