# MATH 223 Python Assignment 1

* Write your solutions in the cells with `YOUR CODE HERE`.
* It's a good idea to copy/paste code from this notebook into a new notebook for rough work and then enter your solutions below. 
* Students are encouraged to collaborate but each student must submit their own notebook.
* Run the test cells to verify your work. **Warning!** There are hidden tests for grading that are only run after you submit your work. Your solutions may not be entirely correct even if they pass all tests below.
* Submit this notebook file (`.ipynb`) to Canvas when complete.

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
```

## Problem 1: Linear Transformations in 2D

### Import Dataset

Import the dataset `points.csv`:

```python
points = np.loadtxt('points.csv',delimiter=',')
```

The dataset `points.csv` is a 2 by 1000 matrix:

```python
points.shape
```

Each column in `points` represents a point in $\mathbb{R}^2$. Create a scatter plot to visualize the points. In the cell below, the parameter `s = 2` specifies the size of the points in the plot and `c = [[0.75,0,0,1]]` specifies the color as an RGBA vector `[red,green,blue,opacity]`.

```python
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Rotation Matrix

A rotation matrix is given by

$$
R_{\theta} = \begin{bmatrix} \cos(\theta) & \sin(\theta) \\ -\sin(\theta) & \cos(\theta) \end{bmatrix}
$$

The linear transformation corresponding to $R_{\theta}$ is rotation clockwise by $\theta$ radians.

```python
theta = np.pi/6
R = np.array([[np.cos(theta),np.sin(theta)],[-np.sin(theta),np.cos(theta)]])
points1 = R@points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points1[0,:],points1[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Scaling Matrix

A scaling matrix is given by

$$
C_{a,b} = \begin{bmatrix} a & 0 \\ 0 & b \end{bmatrix}
$$

The linear transformation corresponding to $C_{a,b}$ is scaling by $a$ is the $x$-direction and scaling by $b$ in the $y$-direction.

```python
a = 3
b = 1/2
C = np.diag([a,b])
points2 = C@points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points2[0,:],points2[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Shearing Matrix

A shear matrix is given by

$$
S_u = \begin{bmatrix} 1 & u \\ 0 & 1 \end{bmatrix}
$$

The linear transformation corresponding to $S_u$ is shearing by factor $u$ in the $x$-direction.

```python
u = 1
S = np.array([[1,u],[0,1]])
points3 = S@points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points3[0,:],points3[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Matrix Multiplication

Combine linear transformations using matrix multiplication.

```python
points4 = S @ C @ R @ points
plt.scatter(points[0,:],points[1,:],s=2,c=[[0.75,0,0,0.1]])
plt.scatter(points4[0,:],points4[1,:],s=2,c=[[0.75,0,0,1]])
plt.axis('equal'), plt.grid(True)
plt.show()
```

### Problem 1a (3 marks)

Find the matrix corresponding to the linear transformation in the image below. The matrix is a combinations of the transformations outlined above. Save the result as `T1a`.

![prob1a.png](prob1a.png)

<!-- grade_id: cell-10761bcd9cc2f1d5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b6f75e319e387cd2, points: 1, grade: true, solution: false -->
```python
"Verify T1a is a NumPy array of size (2,2). (1 mark)"
assert isinstance(T1a,np.ndarray)
assert T1a.shape == (2,2)
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-efbc00f36cfb9ebd, points: 2, grade: true, solution: false -->
```python
"Verify T1a has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(np.round(T1a,1),[[0.7,-0.4],[-0.7, -0.4]]) , "Check entries to one decimal place."
print("Problem 1a Test 2: Success!")
```

### Problem 1b (3 marks)

Find the matrix corresponding to the linear transformation in the image below. The matrix is a combinations of the transformations outlined above. Save the result as `T1b`.

![prob1b.png](prob1b.png)

<!-- grade_id: cell-55e863bb5f4e7004, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-fc129d3ca27caa67, points: 1, grade: true, solution: false -->
```python
"Verify T1b is a NumPy array of size (2,2). (1 mark)"
assert isinstance(T1b,np.ndarray)
assert T1a.shape == (2,2)
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-f1c325264cbd2791, points: 2, grade: true, solution: false -->
```python
"Verify T1b has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(T1b.sum(),1.5) , "Entries should add up to 1.5."
print("Problem 1b Test 2: Success!")
```

## Problem 2: Solving Systems of Equations

Compute the solution of a linear system $A \mathbf{x} = \mathbf{b}$ using the function `scipy.linalg.solve`. For example, the cell below computes the solution for the system:

$$
\left[ \begin{array}{rrr} -2 & 1 & 0 \\ 1 & -2 & 1 \\ 0 & 1 & -2 \end{array} \right]
\begin{bmatrix} x_0 \\ x_1 \\ x_2 \end{bmatrix}
= \begin{bmatrix} 1 \\ 2 \\ 1 \end{bmatrix}
$$

```python
A = np.array([[-2,1,0],[1,-2,1],[0,1,-2]])
b = np.array([1,2,1])
x = la.solve(A,b)
x
```

### Problem 2a (3 marks)

Compute the solution of the system

$$
\left[ \begin{array}{rrr} 0 & -1 & 1 \\ 1 & 2 & 0 \\ 1 & 0 & -1 \end{array} \right]
\begin{bmatrix} x_0 \\ x_1 \\ x_2 \end{bmatrix}
= \begin{bmatrix} 1 \\ 0 \\ 1 \end{bmatrix}
$$

Save the result as `x2a`.

<!-- grade_id: cell-3622ebbb281250ef, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-9e409221661900f6, points: 1, grade: true, solution: false -->
```python
"Verify x2a is a NumPy array of length 3. (1 mark)"
assert isinstance(x2a,np.ndarray)
assert x2a.size == 3
print("Problem 2a Test 1: Success!")
```

<!-- grade_id: cell-53fe853280cbd3dc, points: 2, grade: true, solution: false -->
```python
"Verify x2a has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(x2a.sum(),1.0) , "Entries should add up to 1.0."
print("Problem 1b Test 2: Success!")
```

### Problem 2b (3 marks)

Find the unique polynomial $p(x) = c_0 + c_1 x + c_2 x^2$ which satisfies $p(-1)=3,p(1)=1,p(2)=2$ by solving the linear system

$$
\left[ \begin{array}{rrr} 1 & -1 & \phantom{+}1 \\ 1 & 1 & 1 \\ 1 & 2 & 4 \end{array} \right]
\begin{bmatrix} c_0 \\ c_1 \\ c_2 \end{bmatrix}
= \begin{bmatrix} 3 \\ 1 \\ 2 \end{bmatrix}
$$

Save the vector of coefficients as `c2b`.

<!-- grade_id: cell-52cfbe4f1f0400d8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# Plot the result to verify your solution.
x = np.linspace(-2,3,50)
y = c2b[0] + c2b[1]*x + c2b[2]*x**2
plt.plot(x,y,'b')
plt.plot([-1,1,2],[3,1,2],'r.',ms=10)
plt.show()
```

<!-- grade_id: cell-b498c19c8fca9282, points: 1, grade: true, solution: false -->
```python
"Verify c2b is a NumPy array of length 3. (1 mark)"
assert isinstance(c2b,np.ndarray)
assert c2b.size == 3
print("Problem 2b Test 1: Success!")
```

<!-- grade_id: cell-34c9290b92a10421, points: 2, grade: true, solution: false -->
```python
"Verify c2b has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(c2b.sum(),1.0) , "Entries should add up to 1.0."
print("Problem 2b Test 2: Success!")
```

## Problem 3: Numerical Integration

Fix an integer $N$. Consider the matrix:

$$
M =
\frac{1}{N}
\begin{bmatrix}
0 & 0 & 0 & 0 & \cdots & 0 \\
1 & 0 & 0 & 0 & \cdots & 0 \\
1 & 1 & 0 & 0 & \cdots & 0 \\
1 & 1 & 1 & 0 & \ddots & 0 \\
\vdots & \vdots & \vdots & \ddots & \ddots & \vdots \\
1 & 1 & 1 & \cdots & 1 & 0 \\
\end{bmatrix}
$$

where $M$ has:

* $N+1$ rows and $N+1$ columns
* $1/N$ in all entries below the diagonal
* value 0 on and above the diagonal

### Problem 3a (3 marks)

Use a combination of the functions `np.tril`, `np.ones` and `np.eye` to construct the matrix $M$ above for $N=100$. Save the result as `M3a`. See the [NumPy documentation](https://numpy.org/doc/stable/) to find out how to use the functions `np.tril`, `np.ones` and `np.eye`.

<!-- grade_id: cell-3d8c2476c580d87d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-dda367da8e6edc1a, points: 1, grade: true, solution: false -->
```python
"Verify M3a is a NumPy array of size (101,101). (1 mark)"
assert isinstance(M3a,np.ndarray)
assert M3a.shape == (101,101)
print("Problem 3a Test 1: Success!")
```

<!-- grade_id: cell-2293424461329c09, points: 2, grade: true, solution: false -->
```python
"Verify M3a has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(M3a[1:3,:2],[[0.01,0.],[0.01,0.01]]) , "Entries in rows 1 and 2 and columns 1 and 2 should be [0.01,0.],[0.01,0.01]."
print("Problem 3a Test 2: Success!")
```

### Problem 3b (3 marks)

The matrix $M$ above performs numerical integration (by left Riemann sums) on the vector $\mathbf{y} \in \mathbb{R}^{N+1}$. Let's see why this is the case.

Let $\mathbf{x}$ be the vector with $N+1$ evenly spaced entries from 0 to 1:

$$
\mathbf{x}
= \begin{bmatrix} x_0 \\ x_1 \\ x_2 \\ \vdots \\ x_{N-1} \\ x_N \end{bmatrix}
= \begin{bmatrix} 0 \\ 1/N \\ 2/N \\ \vdots \\ (N-1)/N \\ 1 \end{bmatrix}
$$

In other words, $x_k = k/N$ for $k=0,\dots,N$. Let $f(x)$ be a function on $[0,1]$ and define $\mathbf{y} \in \mathbb{R}^{N+1}$ by $y_k = f(x_k)$ for $k=0,\dots,N$. The vector $\mathbf{w} = M \mathbf{y}$ is an approximation of the function

$$
F(x) = \int_0^x f(t) dt
$$

In particular, if $w_k$ is the entry of $\mathbf{w}$ at index $k$ (starting from index 0) then

$$
w_k = \frac{1}{N} \sum_{i=0}^{k-1} y_i
$$

This is the left Riemann sum for $f(x)$ on the interval $[0,k/N]$.

Use the functions `np.linspace` and `np.exp` to construct the vector $\mathbf{y}$ for $N=100$ and $f(x) = e^{-10x^2}$. Save the result as `y3b`.

<!-- grade_id: cell-3b6c74d0125c60d8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# Plot the result to verify your solution
plt.plot(np.linspace(0,1,101),y3b,'.'), plt.grid(True)
plt.show()
```

<!-- grade_id: cell-db518fda37c06b78, points: 1, grade: true, solution: false -->
```python
"Verify y3b is a NumPy array of length 101. (1 mark)"
assert isinstance(y3b,np.ndarray)
assert y3b.size == 101
print("Problem 3b Test 1: Success!")
```

<!-- grade_id: cell-b3defdddd074c854, points: 2, grade: true, solution: false -->
```python
"Verify y3b has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(y3b[20],0.6703200460356392) , "Entry at index 20 should be 0.6703200460356392."
assert np.allclose(y3b[60],0.02732372244729257) , "Entry at index 60 should be 0.02732372244729257."
print("Problem 3b Test 2: Success!")
```

### Problem 3c (3 marks)

Use matrix multiplication to construct the vector $\mathbf{w} = M \mathbf{y}$ where $M$ and $\mathbf{y}$ are `M3a` and `y3b` respectively. Save the result as `w3c`.

<!-- grade_id: cell-9b9bf29c4595b5ce, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# Plot the result to verify your solution
plt.plot(np.linspace(0,1,101),w3c,'.'), plt.grid(True)
plt.show()
```

<!-- grade_id: cell-7095446c5b198431, points: 1, grade: true, solution: false -->
```python
"Verify w3c is a NumPy array of length 101. (1 mark)"
assert isinstance(w3c,np.ndarray)
assert w3c.size == 101
print("Problem 3c Test 1: Success!")
```

<!-- grade_id: cell-72f06d918fce5c8d, points: 2, grade: true, solution: false -->
```python
"Verify w3c has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.allclose(w3c[20],0.17787686111744105) , "Entry at index 20 should be 0.17787686111744105."
assert np.allclose(w3c[60],0.2830670905645971) , "Entry at index 60 should be 0.2830670905645971."
print("Problem 3c Test 2: Success!")
```

## Problem 4: Gaussian Elimination

The functions below perform the elementary row operations in the Gaussian elimination algorithm.

```python
def add_rows(A,c,i,j):
    "Add c times row i to row j in matrix A."
    Ac = A.copy()
    Ac[j,:] = c*Ac[i,:] + Ac[j,:]
    return Ac

def switch_rows(A,i,j):
    "Interchange row i and row j in matrix A."
    Ac = A.copy()
    Ai = Ac[i,:].copy()
    Aj = Ac[j,:].copy()
    Ac[j,:] = Ai
    Ac[i,:] = Aj
    return Ac

def scale_row(A,c,i):
    "Multiply row i by c in matrix A."
    Ac = A.copy()
    Ac[i,:] = c*Ac[i,:]
    return Ac
```

For example, let's find the reduced row echelon form for the matrix

$$
A = \begin{bmatrix} 2 & 4 & 0 \\ 1 & 2 & 2 \\ 3 & 1 & -2 \end{bmatrix}
$$

Note that we should enter the values as floats (ie. decimals) so that the matrix entries are floats and not integers.

```python
A = np.array([[2.,4.,0.],[1.,2.,2.],[3.,1.,-2.]])
A
```

```python
A1 = add_rows(A,-1/2,0,1)
A1
```

```python
A2 = add_rows(A1,-3/2,0,2)
A2
```

```python
A3 = switch_rows(A2,1,2)
A3
```

```python
A4 = scale_row(A3,1/2,0)
A5 = scale_row(A4,-1/5,1)
A6 = scale_row(A5,1/2,2)
A6
```

### Problem 4a (3 marks)

Apply the Gaussian elimination algorithm to the matrix

$$
A = \begin{bmatrix}
2 & 1 & 0 & 1 \\
1 & 2 & 1 & 0 \\
0 & 1 & 2 & 1 \\
1 & 0 & 1 & 2
\end{bmatrix}
$$

Save the reduced row echelon form of $A$ as `rref4a`.

<!-- grade_id: cell-47929b73fb2a02af, points: 0, grade: false, solution: true -->
```python
A = np.array([[2.,1.,0.,1.],[1.,2.,1.,0.],[0.,1.,2.,1.],[1.,0.,1.,2.]])

# YOUR CODE HERE
```

<!-- grade_id: cell-e7cf9f1a785965b2, points: 1, grade: true, solution: false -->
```python
"Verify rref4a is a NumPy array of size (4,4). (1 mark)"
assert isinstance(rref4a,np.ndarray)
assert rref4a.shape == (4,4)
print("Problem 4a Test 1: Success!")
```

<!-- grade_id: cell-31ed40a43d777053, points: 2, grade: true, solution: false -->
```python
"Verify rref4a has correct entries up to 8 decimal places. This cell contains hidden tests. (2 marks)"
assert np.all(np.abs(np.diag(rref4a)) <= 1.) , "All diagonal entries should be either 0 or 1."
assert np.allclose(rref4a[2:,:2],[[0.,0.],[0.,0.]]) , "All entries below diagonal should be 0."
print("Problem 4a Test 2: Success!")
```

### Problem 4b (1 mark)

Determine the rank of $A$ by inspecting the reduced row echelon form of $A$. Save the result as `r4b`

<!-- grade_id: cell-d0dfb80817e660a7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-81c425075162e770, points: 1, grade: true, solution: false -->
```python
"Verify r4b is the correct value. This cell contains tests. (1 mark)"
assert r4b > 0
print("Problem 4b Test 1: Success!")
```

### Problem 4c (1 mark)

Use the function `np.linalg.matrix_rank` to compute the rank of the matrix `M3a` from Probem 3a above. Use the result to determine the dimension of the nullspace. Save the result (ie. the dimension of the nullspace of `M3a`) as `n4c`. See the [NumPy documentation](https://numpy.org/doc/stable/reference/generated/numpy.linalg.matrix_rank.html).

<!-- grade_id: cell-8b49a39728718d4f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: cell-b4f8dc6778adec1d, points: 1, grade: true, solution: false -->
```python
"Verify n4c is the correct value. This cell contains tests. (1 mark)"
assert n4c > 0
print("Problem 4c Test 1: Success!")
```