# MATH 223 Python Assignment 2

* Write your solutions in the cells with `YOUR CODE HERE`.
* It's a good idea to copy/paste code from this notebook into a new notebook for rough work and then enter your solutions below. 
* Students are encouraged to collaborate but each student must submit their own notebook.
* Run the test cells to verify your work. **Warning!** There are hidden tests for grading that are only run after you submit your work. Your solutions may not be entirely correct even if they pass all tests below.
* Submit this notebook file (`.ipynb`) to Canvas when complete.

```python
import numpy as np
import scipy.linalg as la
import matplotlib.pyplot as plt
```

## How to use `la.eig`

The function `la.eig` computes eigenvalues and eigenvectors of a matrix $A$. For example, consider the matrix

$$
A_0 = \begin{bmatrix} 1 & 2 \\ 2 & 1 \end{bmatrix}
$$

The eigenvalues are $\lambda_0 = 3$ and $\lambda_1 = -1$ with corresponding unit eigenvectors

$$
\vec{v}_0 = \frac{1}{\sqrt{2}} \begin{bmatrix} 1 \\ 1 \end{bmatrix}
\hspace{5mm}
\vec{v}_1 = \frac{1}{\sqrt{2}} \begin{bmatrix} -1 \\ 1 \end{bmatrix}
$$

Let's use `la.eig` to compute $\lambda_0$, $\lambda_1$, $\vec{v}_0$, and $\vec{v}_1$.

```python
A0 = np.array([[1,2],[2,1]])
D,V = la.eig(A0)
```

```python
D
```

```python
V
```

The function `la.eig` returns a vector `D` of eigenvalues and a matrix `V` with unit eigenvectors in the columns. In particular, eigenvalue `D[0]` corresponds to eigenvector `V[:,0]`, `D[1]` corresponds to `V[:,1]`, etc.

Also note that the vector `D` of eigenvalues is always complex even if the eigenvalues are real numbers. We can convert the values to real numbers using the function `np.real`.

```python
D = np.real(D)
```

```python
D[0]
```

```python
V[:,0]
```

Compare the eigenvector we comptued by hand:

$$
\vec{v}_0 = \frac{1}{\sqrt{2}} \begin{bmatrix} 1 \\ 1 \end{bmatrix}
$$

```python
(1/2**0.5)*np.array([1,1])
```

They are the same!

## Problem 1

### Problem 1a (2 marks)

Use the functions `np.arange` (see [documentation](https://numpy.org/doc/stable/reference/generated/numpy.arange.html)) and `np.reshape` (see [documentation](https://numpy.org/doc/stable/reference/generated/numpy.reshape.html)) to create the $10 \times 10$ matrix with entries from 0 to 99 which increase across the rows:

$$
A_1 = \begin{bmatrix}
0 & 1 & 2 & \cdots & 9 \\
10 & 11 & 12 & \cdots & 19 \\
\vdots & \vdots & \vdots & \ddots & \vdots \\
90 & 91 & 92 & \cdots & 99
\end{bmatrix}
$$

Save the result as `A1`.

<!-- grade_id: cell-93c24c8725fe4036, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A1 = 
```

<!-- grade_id: cell-f1d568efffcaf47e, points: 1, grade: true, solution: false -->
```python
"Verify A1 is a NumPy array of size (10,10). (1 mark)"
assert isinstance(A1,np.ndarray)
assert A1.shape == (10,10)
print("Problem 1a Test 1: Success!")
```

<!-- grade_id: cell-74ac8492d697a9e0, points: 1, grade: true, solution: false -->
```python
"Verify A1 has correct values. This cell contains hidden tests. (1 mark)"
assert A1[0,0] == 0 and A1[9,9] == 99
print("Problem 1a Test 2: Success!")
```

### Problem 1b (2 marks)

The eigenvalues of $A_1$ are real numbers. Find the largest eigenvalue of $A_1$ and save the result as `lambda1` (as a real number).

<!-- grade_id: cell-5565c2f4d4d7b8d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
lambda1 = 
```

<!-- grade_id: cell-6477756710098b0a, points: 1, grade: true, solution: false -->
```python
"Verify lambda1 is a positive number. (1 mark)"
assert lambda1 > 0
print("Problem 1b Test 1: Success!")
```

<!-- grade_id: cell-0d0c08c2458e112b, points: 1, grade: true, solution: false -->
```python
"Verify value of lambda1. This cell contains hidden tests. (1 mark)"
assert lambda1 > 500
print("Problem 1b Test 2: Success!")
```

### Problem 1c (2 marks)

Determine the algebraic multiplicity of the eigenvalue $\lambda=0$ for $A_1$. Save the result as `algmult`.

<!-- grade_id: cell-e5ee8def2179df0c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
algmult = 
```

<!-- grade_id: cell-af4b363107b1e780, points: 1, grade: true, solution: false -->
```python
"Verify algmult is a positive number. (1 mark)"
assert algmult > 0
print("Problem 1c Test 1: Success!")
```

<!-- grade_id: cell-797b15add819405f, points: 1, grade: true, solution: false -->
```python
"Verify value of algmult. This cell contains hidden tests. (1 mark)"
assert algmult > 0
print("Problem 1c Test 2: Success!")
```

## Problem 2 (3 marks)

Construct a $3 \times 3$ matrix $A_2$ with eigenvalues $\lambda_1 = -1$, $\lambda_2 = 1$, $\lambda_3 = 2$ such that all entries of $A_2$ are nonzero. Save the result as `A2`.

Hint: create $A_2 = PDP^{-1}$ for matrices $P$ and $D$ using the function `la.inv` (see [documentation](https://docs.scipy.org/doc/scipy/reference/generated/scipy.linalg.inv.html)).

<!-- grade_id: cell-799a6f001382c209, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A2 = 
```

<!-- grade_id: cell-8e23091aa2aa9f98, points: 1, grade: true, solution: false -->
```python
"Verify A2 is a NumPy array of size (3,3). (1 mark)"
assert isinstance(A2,np.ndarray)
assert A2.shape == (3,3)
print("Problem 2 Test 1: Success!")
```

<!-- grade_id: cell-74a8df330cc992a2, points: 1, grade: true, solution: false -->
```python
"Verify entries of A2 are nonzero real numbers. (1 mark)"
assert np.all(A2 != 0)
assert np.max(np.abs(A2.imag)) == 0.
print("Problem 2 Test 2: Success!")
```

<!-- grade_id: cell-f94cfbc568e59788, points: 1, grade: true, solution: false -->
```python
"Verify eigenvalues of A2 are -1,1,2. (1 mark)"
D2,_ = la.eig(A2)
D2 = np.sort(np.real(D2))
assert np.allclose(D2,[-1,1,2])
print("Problem 2 Test 3: Success!")
```

## Problem 3

Consider the matrix $A_3$ of size $100 \times 100$ with $2$ in the diagonal and $1$ in the upper and lower diagonals

$$
A_3 = \left[ \begin{array}{rrrrr} 2 & 1 & & & \\ 1 & 2 & 1 & & \\ & \ddots & \ddots & \ddots & \\ & & 1 & 2 & 1 \\ & & & 1 & 2 \end{array} \right]
$$

### Problem 3a (2 marks)

Use the functions `np.eye` (see [documentation](https://numpy.org/devdocs/reference/generated/numpy.eye.html)) and `np.diag` (see [documentation](https://numpy.org/devdocs/reference/generated/numpy.diag.html)) to construct the matrix $A_3$ and save the result as `A3`.

<!-- grade_id: cell-7d75179b87543fc7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A3 = 
```

<!-- grade_id: cell-825c1ec1229044fb, points: 1, grade: true, solution: false -->
```python
"Verify A3 is a NumPy array of size (100,100). (1 mark)"
assert isinstance(A3,np.ndarray)
assert A3.shape == (100,100)
print("Problem 3a Test 1: Success!")
```

<!-- grade_id: cell-76a5c72225cd5651, points: 1, grade: true, solution: false -->
```python
"Verify A3 has correct values. This cell contains hidden tests. (1 mark)"
assert np.allclose(A3[0,:3],[2,1,0])
assert np.allclose(A3[-1,-3:],[0,1,2])
print("Problem 3a Test 2: Success!")
```

### Problem 3b (2 marks)

Use the functions `la.eig`, `np.sort` (see [documentation](https://numpy.org/devdocs/reference/generated/numpy.sort.html)) and `np.real` (see [documentation](https://numpy.org/devdocs/reference/generated/numpy.real.html)) to compute the eigenvalues of $A_3$ and sort the (real) eigenvalues into increasing order. Save the result as `A3_evals`.

<!-- grade_id: cell-be8c6cfb17e03a26, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A3_evals = 
```

<!-- grade_id: cell-446d0a73e4eda58e, points: 1, grade: true, solution: false -->
```python
"Verify A3_evals is a NumPy array of length (100). (1 mark)"
assert isinstance(A3_evals,np.ndarray)
assert A3_evals.size == 100
print("Problem 3b Test 1: Success!")
```

<!-- grade_id: cell-40eabb819cbf1a06, points: 1, grade: true, solution: false -->
```python
"Verify A3_evals has correct values. This cell contains hidden tests. (1 mark)"
assert np.allclose(A3_evals[:2],[0.00096744, 0.00386881])
assert np.allclose(A3_evals[-2:],[3.99613119, 3.99903256])
print("Problem 3b Test 2: Success!")
```

If done correctly, the cell below creates the figure below:

```python
plt.plot(A3_evals,'.')
plt.plot()
plt.show()
```

![Problem 3a](https://personal.math.ubc.ca/~pwalls/data/A3_evals.png)

### Proboem 3c (3 marks)

Let $\lambda_0,\lambda_1,\dots,\lambda_{99}$ be the eigenvalues of $A_3$ in increasing order. Find integers $a$, $b$ and $c$ such that $\lambda_n = a + b \cos(\pi n / c)$. Save the values as `a3`, `b3` and `c3`. Note that you can determine the integers $a$, $b$ and $c$ simply by inspection. This problem shows that the eigenvalues of $A_3$ satiafy the beautifu formula:

$$
\lambda_n = a + b \cos(\pi n / c)
$$

<!-- grade_id: cell-dd840d82aae63ef6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# Verify the result
N = 100
n = np.arange(N)
plt.plot(n,a3 + b3*np.cos(np.pi*n/c3),'b',alpha=0.5)
plt.plot(n,A3_evals,'b.',alpha=0.5)
plt.show()
```

<!-- grade_id: cell-cc2430985bae2603, points: 1, grade: true, solution: false -->
```python
"Verify a3. This cell contains hidden tests. (1 mark)"
assert a3 > 0
print("Problem 3c Test 1: Success!")
```

<!-- grade_id: cell-480720d783750bfb, points: 1, grade: true, solution: false -->
```python
"Verify b3. This cell contains hidden tests. (1 mark)"
assert b3 < 0
print("Problem 3c Test 2: Success!")
```

<!-- grade_id: cell-403845108624c273, points: 1, grade: true, solution: false -->
```python
"Verify c3. This cell contains hidden tests. (1 mark)"
assert c3 > 0
print("Problem 3c Test 3: Success!")
```

## Problem 4

Consider the matrix $A_4$ with 0 on the diagonal and $1/(N-1)$ everywhere else:

$$
A_4 = \frac{1}{N-1} \begin{bmatrix}
0 & 1 & 1 & \cdots & 1 \\
1 & 0 & 1 & \cdots & 1 \\
1 & 1 & 0 & \cdots & 1 \\
\vdots & \vdots & \vdots & \ddots & \vdots \\
1 & 1 & 1 & \cdots & 0 \\
\end{bmatrix}
$$

### Problem 4a (2 marks)

Use the functions `np.eye` and `np.ones` (see [documentation](https://numpy.org/doc/stable/reference/generated/numpy.ones.html) to create the matrix $A_4$ for $N = 200$. Save the result as $A4$.

<!-- grade_id: cell-13fe6cbed2afbf56, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A4 = 
```

<!-- grade_id: cell-4eb0e1297bd436f7, points: 1, grade: true, solution: false -->
```python
"Verify A4 is a NumPy array of size (200,200). (1 mark)"
assert isinstance(A4,np.ndarray)
assert A4.shape == (200,200)
print("Problem 4a Test 1: Success!")
```

<!-- grade_id: cell-6cc9d7cb40fe277f, points: 1, grade: true, solution: false -->
```python
"Verify A4 has correct values. This cell contains hidden tests. (1 mark)"
assert A4[0,0] == 0
assert np.abs(A4[1,0] - 1/199) < 1e-10
print("Problem 4a Test 2: Success!")
```

### Problem 4b (2 marks)

Determine the matrix $A_{\infty}$ in the limit $A_{\infty} = \lim_{k \to \infty} A^k$. Save the result as `A4infty`.

<!-- grade_id: cell-a3b9e2aeb054ce8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A4infty = 
```

<!-- grade_id: cell-4b692ccb931c25c5, points: 1, grade: true, solution: false -->
```python
"Verify A4infty is a NumPy array of size (200,200). (1 mark)"
assert isinstance(A4infty,np.ndarray)
assert A4infty.shape == (200,200)
print("Problem 4b Test 1: Success!")
```

<!-- grade_id: cell-b0c6fa0d146be3cd, points: 1, grade: true, solution: false -->
```python
"Verify A4 has correct values. This cell contains hidden tests. (1 mark)"
assert np.abs(la.det(A4infty)) < 1e-10
print("Problem 4b Test 2: Success!")
```

## Problem 5 (2 marks)

The matrix $R_5$ below is a 3D rotation matrix:

```python
R5 = np.array([[-0.47492144, -0.22606924,  0.85049534],
               [ 0.81603782, -0.47492144,  0.32944181],
               [ 0.32944181,  0.85049534,  0.41003142]])
```

Plot the cube with vertices at $(\pm 1, \pm 1, \pm 1)$ and plot the rotation by $R_5$:

```python
ax = plt.figure().add_subplot(projection='3d')

cube = np.array([[-1,1,1,-1,-1,-1,-1,1,1,-1,-1,1,1,1,1,-1,-1],
                 [-1,-1,-1,-1,-1,1,1,1,1,1,1,1,-1,-1,1,1,-1],
                 [-1,-1,1,1,-1,-1,-1,-1,1,1,-1,-1,-1,1,1,1,1]])
ax.plot(cube[0,:],cube[1,:],cube[2,:],'b');

cubeR5 = R5@cube
ax.plot(cubeR5[0,:],cubeR5[1,:],cubeR5[2,:],'r',alpha=0.7);

ax.set_xlim([-2,2]), ax.set_ylim([-2,2]), ax.set_zlim([-2,2])
ax.set_xlabel('x'); ax.set_ylabel('y'); ax.set_zlabel('z');
plt.show()
```

Find the axis of rotation of $R_5$ and save the result as `v5`. (The vector should be **unit** vector with **real** entries such that the `z` entry is positive so that the axis is pointing up when plotted in $\mathbb{R}^3$.)

<!-- grade_id: cell-d992624fe6900e02, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
v5 = 
```

<!-- grade_id: cell-f067f30c796cf2a9, points: 1, grade: true, solution: false -->
```python
"Verify v5 is a NumPy array of length 3. (1 mark)"
assert isinstance(v5,np.ndarray)
assert v5.size == 3
print("Problem 5 Test 1: Success!")
```

<!-- grade_id: cell-e73fab9a83e812a8, points: 1, grade: true, solution: false -->
```python
"Verify v5 has correct values. This cell contains hidden tests. (1 mark)"
assert np.abs(v5[0] - 1/np.sqrt(6)) < 1e-8
assert v5[2] > 0
print("Problem 5 Test 2: Success!")
```