# MATH 215 Python Assignment 3

## Instructions

* Enter your solutions in cells with `YOUR CODE HERE` (all other cells are ignored by the autograder)
* Run the tests to verify your work (and note some cells contain hidden tests)
* Hidden tests only run during autograding
* Do not import any pacakges (other than the standard packages in the cell below)
* See [Python for UBC Math](https://ubcmath.github.io/python/) for instructions and examples
* Submit the completed notebook to Canvas (download `.ipynb` file to your machine and then upload to Canvas)

```python
import numpy as np
import matplotlib.pyplot as plt
```

## Python Functions

The Python package `numpy` contains all the usual mathematical functions such as `np.exp`, `np.cos` and `np.sin`. Use the `lambda` keyword to create new Python functions.

For example, the code below creates a function `f` which takes input parameter `x` and returns the value of the function

$$
f(x) = e^{-0.1x} \cos(\pi x)
$$

```python
f = lambda x: np.exp(-0.1*x)*np.cos(np.pi*x)
```

```python
f(0)
```

```python
f(1)
```

```python
f(5)
```

## Plotting Functions

Use the function `f` to plot the graph $y = f(x)$. We construct a vector `x` of 1000 equally spaced values from $x=0$ to $x=50$ and then apply `f` to `x` to generate the vector of $y$ values. The number 1000 is arbitrary and we simply choose it to be large enough to plot the function smoothly. The function `plt.plot` simply connects the dots defined by the coordinates in vectors `x` and `y` and so we need lots of points to plot this particular function smoothly since it has a lot of oscillation.

```python
x = np.linspace(0,50,1000)
y = f(x)
plt.plot(x,y)
plt.show()
```

See [Python for UBC Math > Plotting](https://ubcmath.github.io/python/getting-started/vectorization.html) for more details and examples.

## Problem 1 (5 marks)

Find the unique solution of $y'' + 3y' + 2y = 0$, $y(0) = 1$, $y'(0) = 0$. Save the solution as a Python function called `y1`.

<!-- grade_id: cell-f3550bb3a969af2f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y1 = lambda t: 
```

<!-- grade_id: cell-07206cfc41278d68, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if y1 is defined as a function. (1 mark)
assert callable(y1) , "y1 should be a function."
print("Test 1: Success!")
```

<!-- grade_id: cell-99d8946ff0e80da7, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if y1(0) = 1. (1 mark)
assert np.allclose(y1(0),1) , "y1(0) should be 1."
print("Test 2: Success!")
```

<!-- grade_id: cell-d1c831dbffa73b8d, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if y1 is near 0 for large t. (1 mark)
assert np.allclose(y1(100),0) , "y1(100) should be near 0."
print("Test 3: Success!")
```

<!-- grade_id: cell-e7efea905f533e9e, points: 1, grade: true, solution: false -->
```python
# Test 4: Check if y1 returns correct value. (1 mark)
assert np.allclose(y1(0.5),0.8451818782538245) , "y1(0.5) should be 0.8451818782538245."
print("Test 4: Success!")
```

<!-- grade_id: cell-560f3324945fe6f2, points: 1, grade: true, solution: false -->
```python
# Test 5: Check if y1 returns correct values. This cell contains hidden tests. (1 mark)
```

Plot the function `y1` over the interval $[0,5]$.

```python
t1 = np.linspace(0,5,200)
plt.plot(t1,y1(t1))
plt.show()
```

## Problem 2 (5 marks)

Find the unique solution of $y'' + 2y' + 3y = 0$, $y(0) = 1$, $y'(0) = 0$. Save the solution as a Python function called `y2`.

<!-- grade_id: cell-8f82f26a74ac42a3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y2 = lambda t: 
```

<!-- grade_id: cell-2b91584de42ab53b, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if y2 is defined as a function. (1 mark)
assert callable(y2) , "y2 should be a function."
print("Test 1: Success!")
```

<!-- grade_id: cell-94cdb841f9d20e69, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if y2(0) = 1. (1 mark)
assert np.allclose(y2(0),1) , "y2(0) should be 1."
print("Test 2: Success!")
```

<!-- grade_id: cell-7bdc6cb0018a668c, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if y2 is near 0 for large t. (1 mark)
assert np.allclose(y2(100),0) , "y2(100) should be near 0."
print("Test 3: Success!")
```

<!-- grade_id: cell-0d60c67f41b61c4a, points: 1, grade: true, solution: false -->
```python
# Test 4: Check if y2 returns correct value. (1 mark)
assert np.allclose(y2(1),0.314316081745773) , "y2(1) should be 0.314316081745773."
print("Test 4: Success!")
```

<!-- grade_id: cell-eda1b68751ec8f44, points: 1, grade: true, solution: false -->
```python
# Test 5: Check if y2 returns correct values. This cell contains hidden tests. (1 mark)
```

Plot the function `y2` over the interval $[0,5]$.

```python
t2 = np.linspace(0,5,200)
plt.plot(t2,y2(t2))
plt.show()
```

### Problem 3 (7 marks)

Consider a mass-spring system with damping

$$
mx'' + cx' + kx = F_0 \cos(\omega t)
$$

The steady state response of the system (ie. the particular solution) is

$$
x_{ss}(t) = A \cos(\omega t - \gamma)
$$

where the amplitude is given by

$$
A = \frac{F_0}{m \sqrt{(2 \omega p)^2 + (\omega_0^2 - \omega^2)^2 }}
$$

$$
p = \frac{c}{2m} \ , \ \ \omega_0 = \sqrt{\frac{k}{m}} \ , \ \ \gamma = \arctan\left( \frac{2 \omega p}{\omega_0^2 - \omega^2}\right)
$$

If $\omega_0^2 - 2p^2 > 0$, then the maximum amplitude occurs when the forcing frequency is

$$
\omega = \sqrt{\omega_0^2 - 2p^2}
$$

This is known as the practical resonant frequency. See [DiffyQs 2.6.2](https://www.jirka.org/diffyqs/html/forcedo_section.html).

Let $m=12$, $c = 4$, $k = 80$ and $F_0 = 100$. Create a Python function called `A` which takes an input parameter `w` (which represents the forcing frequency $\omega$) and returns the amplitude $A$. Compute the practical resonant frequency and save the result as `wmax`.

<!-- grade_id: cell-51bbd162c23863fd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
A = lambda w: 
wmax = 
```

<!-- grade_id: cell-a4277b37c58f94dd, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if A is defined as a function. (1 mark)
assert callable(A) , "A should be a function."
print("Test 1: Success!")
```

<!-- grade_id: cell-a295dc2eb2acca5b, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if A(0) = F0/k. (1 mark)
assert np.allclose(A(0),1.25) , "A(0) should be 1.25."
print("Test 2: Success!")
```

<!-- grade_id: cell-323f2019d171eaa8, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if A is near 0 for large w. (1 mark)
assert np.abs(A(10000)) < 1e-6 , "A(10000) should be near 0."
print("Test 3: Success!")
```

<!-- grade_id: cell-a176217d7839e737, points: 1, grade: true, solution: false -->
```python
# Test 4: Check if A returns correct value. (1 mark)
assert np.allclose(A(1),1.468050548786759) , "A(1) should be 1.468050548786759."
print("Test 4: Success!")
```

<!-- grade_id: cell-2c8fc4f0be694350, points: 1, grade: true, solution: false -->
```python
# Test 5: Check if A returns correct values. This cell contains hidden tests. (1 mark)
```

<!-- grade_id: cell-aa8fb6a14e62fd04, points: 1, grade: true, solution: false -->
```python
# Test 6: Check if wmax is between 2 and 3. (1 mark)
assert wmax > 2 , "wmax should be larger than 2."
assert wmax < 3 , "wmax should be less then 3."
print("Test 6: Success!")
```

<!-- grade_id: cell-b1b9ebb2d45bd115, points: 1, grade: true, solution: false -->
```python
# Test 7: Check if wmax is the correct value. This cell contains hidden tests. (1 mark)
```

Plot the amplitude $A$ of the steady state response as a function of $\omega$ over the interval $0 \leq \omega \leq 8$, and plot the point corresponding to the amplitude of the practical resonant frequency.

```python
w3 = np.linspace(0,8,500)
plt.plot(w3,A(w3))
plt.plot(wmax,A(wmax),'r*')
plt.grid(True)
plt.show()
```

## Problem 4 (5 marks)

Consider a mass-spring system without damping

$$
m x'' + k x = F_0 \cos(\omega t)
$$

The general solution is

$$
x(t) = C_1 \cos(\omega_0 t) + C_2 \sin(\omega_0 t) + \frac{F_0}{m (\omega_0^2 - \omega^2)} \cos(\omega t)
$$

where

$$
\omega_0 = \sqrt{\frac{k}{m}}
$$

is the natural frequency. See [DiffyQs 2.6.1](https://www.jirka.org/diffyqs/html/forcedo_section.html).

Let $m=2$, $k=10$ and $F_0 = 1$ with initial conditions $x(0)=x'(0)=0$. Find the solution $x(t,\omega)$ as a function of both time $t$ and the forcing frequency $\omega$ and save the solution as a Python function called `x4` with input parameters `t` and `w`.

<!-- grade_id: cell-27b0f8e722953174, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
x4 = lambda t,w: 
```

<!-- grade_id: cell-1b2986bfab7b7199, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if x4 is defined as a function. (1 mark)
assert callable(x4) , "x4 should be a function."
print("Test 1: Success!")
```

<!-- grade_id: cell-5d441daf1fb7b28c, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if x4(0,2) = 0. (1 mark)
assert np.allclose(x4(0,2),0) , "x4(0,2) should be 0."
print("Test 2: Success!")
```

<!-- grade_id: cell-f1312be20ff281e4, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if x4 returns correct value. (1 mark)
assert np.allclose(x4(1,1),0.14469689779066328) , "x4(1,1) should be 0.14469689779066328."
print("Test 3: Success!")
```

<!-- grade_id: cell-c8241d79353c0a48, points: 2, grade: true, solution: false -->
```python
# Test 4: Check if x4 returns correct values. This cell contains hidden tests. (2 marks)
```

Plot the function `x4` for $\omega = 1,2,3$ over the interval $0 \leq t \leq 50$.

```python
t4 = np.linspace(0,50,1000)
plt.plot(t4,x4(t4,1),label='w=1')
plt.plot(t4,x4(t4,2),label='w=2')
plt.plot(t4,x4(t4,3),label='w=3')
plt.legend(), plt.grid(True)
plt.show()
```

When the forcing frequency $\omega$ is nearly equal to the natural frequency $\omega_0$, the system displays "beats" where the complementary solution and particular solution produce both constructive and destructive interference. The natural frequency is $\omega_0 = \sqrt{5} \approx 2.236$ and we see "beats" for $\omega = 2$ in the example above.