# MATH 215 Python Assignment 1

## Instructions

* Enter your solutions in cells with `YOUR CODE HERE` (all other cells are ignored by the autograder)
* Run the tests to verify your work
* Some cells contain hidden tests for grading
* Do not import any pacakges (other than the standard packages in the cell below)
* See [Python for UBC Math](https://ubcmath.github.io/python/) for instructions and examples
* Submit the completed notebook to Canvas (download `.ipynb` file to your machine and then upload to Canvas)

```python
import numpy as np
import matplotlib.pyplot as plt
```

## Plotting Solutions and Slope Fields

The slope field of the equation $y' = f(t,y)$ plots a short line of slope $f(t,y)$ at each point in a grid in the $ty$-plane. The function `slopefield` defined below plots the slope field for the function `f(t,y)` over the intervals `tinterval` and `yinterval` with step sizes `tstep` and `ystep`. See [Python for UBC Math](https://ubcmath.github.io/python/differential-equations/slope-fields.html) for more details and examples about plotting functions and slope fields.

```python
def slopefield(f,tinterval,yinterval,tstep,ystep):
    t = np.arange(tinterval[0],tinterval[1],tstep)
    y = np.arange(yinterval[0],yinterval[1],ystep)
    L = 0.7*min(tstep,ystep)
    for i in range(len(t)):
        for j in range(len(y)):
            slope = f(t[i],y[j])
            theta = np.arctan(slope)
            dy = L*np.sin(theta)
            dt = L*np.cos(theta)
            plt.plot([t[i],t[i] + dt],[y[j],y[j] + dy],'b')
```

We don't need to know *exactly* how the function `slopefield` works or what it is doing. We just need to know how to use it.

For example, let's plot the slope field for the equation $y' = -ty$ over the intervals $0 \leq t \leq 4$ and $-2 \leq y \leq 2$ with step size $0.2$ in both directions.

```python
f = lambda t,y: -t*y
slopefield(f,[0,4],[-2,2],0.2,0.2)
```

Use separation of variables to find the general solution $y(t) = C e^{-t^2/2}$. Plot the unique solution $y(t) = e^{-t^2/2}$ which satisfies $y(0)=1$ along with the slope field. The first line in the cell below defines the function $y(t)$ as a Python function called `y`.

```python
y = lambda t: np.exp(-t**2/2)
t = np.linspace(0,4,200)
plt.plot(t,y(t),'r')
f = lambda t,y: -t*y
slopefield(f,[0,4],[-2,2],0.2,0.2)
```

**What's the point?** The slope field shows the shape of all solutions of the differential equation $y' = f(t,y)$. We can verify our analytical solutions by plotting particular solutions in the slope field and observing that they agree. We can also make observations about solutions of equations without finding solutions by simply observing the slope field. See [Python for UBC Math](https://ubcmath.github.io/python/differential-equations/slope-fields.html) for more examples and details.

## Problem 1 (4 marks)

Find the unique solution of the equation $y' = y \cos(\pi t)$, $y(0) = 1$. Enter your solution as a Python function called `y1`.

<!-- grade_id: cell-19bb70bfd66a9976, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
y1 = lambda t: # Enter formula for y1(t)
```

<!-- grade_id: cell-0a7739570dd4c34c, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if y1 is defined as a function. (1 mark)
assert callable(y1) , "y1 should be a Python function."
print("Problem 1 Test 1: Success!")
```

<!-- grade_id: cell-8be89b8bd1358b3d, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if y1(0) is equal to 1. (1 mark)
assert np.allclose(y1(0),1) , "The value y1(0) should be 1."
print("Problem 1 Test 2: Success!")
```

<!-- grade_id: cell-838adbe1a5009e6d, points: 2, grade: true, solution: false -->
```python
# Test 3: This cell contains hidden tests for grading. (2 marks)
```

```python
# Plot the function y1 with the slope field to verify.
t = np.linspace(0,3,200)
plt.plot(t,y1(t),'r')
f = lambda t,y: y*np.cos(np.pi*t)
slopefield(f,[0,3],[0,2],0.1,0.1)
```

## Problem 2 (3 marks)

Consider the first order differential equation $y' = \cos(t) - \sin(y)$. Define the right hand side of the equation as a Python function called `f2` and plot the slope field over the intervals $0 \leq t \leq 5$ and $0 \leq y \leq 5$ with step size $h=0.2$ in both directions.

<!-- grade_id: cell-f7b9bf1a6d76fecc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
f2 = 
```

<!-- grade_id: cell-da2e998ee96c74ce, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if f2 is defined as a function. (1 mark)
assert callable(f2) , "f2 should be a Python function."
print("Problem 2 Test 1: Success!")
```

<!-- grade_id: cell-a03f89545bd8a5ba, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if f2(0,0) is equal to 1, and if f2(pi/2,pi/2) is equal to -1. (1 mark)
assert np.allclose(f2(0,0),1) , "The value f2(0,0) should be 1."
assert np.allclose(f2(np.pi/2,np.pi/2),-1) , "The value f2(pi/2,pi/2) should be -1."
print("Problem 2 Test 2: Success!")
```

<!-- grade_id: cell-cb36cfe408992607, points: 1, grade: true, solution: false -->
```python
# Test 3: This cell contains hidden tests for grading. (1 mark)
```

## Problem 3 (3 marks)

Consider again the equation $y' = \cos(t) - \sin(y)$ in the previous problem. Note that the equation is neither separable nor linear and so we do not have any method to find solutions analytically. Plot the slope field over different intervals and then determine if the statements below are True or False.

(a) If $y(t)$ is any solution of $y' = \cos(t) - \sin(y)$ such that $y(0) > 0$ then $y(t) > 0$ for all $t > 0$. Enter your response as `True` or `False` and save as `observation3a`.

(b) If $y(t)$ is any solution of $y' = \cos(t) - \sin(y)$ such that $y(5) > 0$ then $y(0) > 0$. Enter your response as `True` or `False` and save as `observation3b`.

<!-- grade_id: cell-80b2de1f76bbceb6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
observation3a = 
observation3b = 
```

<!-- grade_id: cell-bc40936164dac6fd, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if observation3a is True or False. (1 mark)
assert isinstance(observation3a,bool) , "observation3a should be either True or False."
assert isinstance(observation3b,bool) , "observation3b should be either True or False."
print("Problem 3 Test 1: Success!")
```

<!-- grade_id: cell-1c9cf1929c785496, points: 1, grade: true, solution: false -->
```python
# Test 2: This cell contains hidden tests for grading. (1 mark)
```

<!-- grade_id: cell-6309009e8fe67716, points: 1, grade: true, solution: false -->
```python
# Test 3: This cell contains hidden tests for grading. (1 mark)
```

## Problem 4 (4 marks)

Find the unique solution of the equation $y' = 1 - y^2$, $y(0)=0$. Plot the solution over the interval $0 \leq t \leq 3$.

In particular, create a vector of $200$ equally spaced $t$ values from $t=0$ to $t=3$. Save the vector as `t4`.

Create the vector of corresponding $y$ values of the unique solution. Save the vector as `y4`.

<!-- grade_id: cell-d195417d94d6dcc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
t4 = 
y4 = 
```

<!-- grade_id: cell-7e7c5ce3f4d2848d, points: 1, grade: true, solution: false -->
```python
# Test 1: Check if t4 and y4 are NumPy arrays. (1 mark)
assert isinstance(t4,np.ndarray) , "t4 should be a NumPy array constructed by np.linspace."
assert isinstance(y4,np.ndarray) , "y4 should be a NumPy array."
print("Problem 4 Test 1: Success!")
```

<!-- grade_id: cell-0ca3619c7d0e4a9b, points: 1, grade: true, solution: false -->
```python
# Test 2: Check if t4 has the correct values. (1 mark)
assert len(t4) == 200 , "t4 should have length 200."
assert np.allclose(t4[0],0) , "First entry of t4 should be 0."
assert np.allclose(t4[-1],3) , "Last entry of t4 should be 3."
print("Problem 4 Test 2: Success!")
```

<!-- grade_id: cell-ad56af915a54542f, points: 1, grade: true, solution: false -->
```python
# Test 3: Check if y4 has the correct values. (1 mark)
assert len(y4) == 200 , "y4 should have length 200."
assert np.allclose(y4[0],0) , "First entry of y4 should be 0."
assert np.allclose(y4[-1],0.9950547536867305) , "Last entry of y4 should be 0.9950547536867305."
print("Problem 4 Test 3: Success!")
```

<!-- grade_id: cell-a631d7b6ed6cf7e8, points: 1, grade: true, solution: false -->
```python
# Test 4: This cell contains hidden tests for grading. (1 marks)
```

```python
# Plot the values y4 with the slope field to verify.
plt.plot(t4,y4,'r')
f = lambda t,y: 1 - y**2
slopefield(f,[0,3],[-2,2],0.2,0.2)
```